```python
# Initialize Otter
import otter
grader = otter.Notebook("hw11.ipynb")
```

<img style="display: block; margin-left: auto; margin-right: auto" src="./ccsf-logo.png" width="250rem;" alt="The CCSF black and white logo">

# Homework 11: Correlation and Regression

## References
* [The Regression Line](https://www.inferentialthinking.com/chapters/15/2/Regression_Line.html)
* [Method of Least Squares](https://www.inferentialthinking.com/chapters/15/3/Method_of_Least_Squares.html)
* [Least Squares Regression](https://www.inferentialthinking.com/chapters/15/4/Least_Squares_Regression.html)
* [`datascience` Documentation](https://datascience.readthedocs.io/)
* [Python Quick Reference](https://ccsf-math-108.github.io/materials-sp24/resources/quick-reference.html)

## Assignment Reminders

- Make sure to run the code cell at the top of this notebook that starts with `# Initialize Otter` to load the auto-grader.
- For all tasks indicated with a 🔎 that you must write explanations and sentences for, provide your answer in the designated space.
- Throughout this assignment and all future ones, please be sure to not re-assign variables throughout the notebook! _For example, if you use `max_temperature` in your answer to one question, do not reassign it later on. Otherwise, you will fail tests that you thought you were passing previously!_
- We encourage you to discuss this assignment with others but make sure to write and submit your own code. Refer to the syllabus to learn more about how to learn cooperatively.
- Unless you are asked otherwise, use the non-interactive visualizations when asked to produce a visualization for a task.
- View the related <a href="https://ccsf.instructure.com" target="_blank">Canvas</a> Assignment page for additional details.

Run the following code cell to import the tools for this assignment.

```python
from datascience import *
import numpy as np
%matplotlib inline
import matplotlib.pyplot as plt
plt.style.use('fivethirtyeight')
```

## Linear Regression Setup

When performing linear regression in this course, you will need to compute several important quantities which will be used throughout our analysis. Unless otherwise specified when asked to make a prediction please assume we are predicting $y$ from $x$ throughout this assignment. To help with our later analysis, we will begin by writing some of these functions and understanding what they can do for us.

### Task 01 📍

Define a function `standard_units` that converts a given array to standard units.

**Hint:** You may find the `np.mean` and `np.std` functions helpful.

_Points:_ 2

```python
def standard_units(data):
    ...
```

```python
grader.check("task_01")
```

### Task 02 📍

Which of the following are true about standard units? Assume we have converted an array of data into standard units using the function above.

1. The unit of all our data when converted into standard units is the same as the unit of the original data.
2. The sum of all our data when converted into standard units is 0.
3. The standard deviation of all our data when converted into standard units is 1.
4. Adding 5 to our original data has no impact on the resultant data when converted to standard units.
5. Multiplying our original data by 5 has no impact on the resultant data when converted to standard units.

Assign `standard_array` to an array of your selections.

_Points:_ 2

```python
standard_array = make_array(...)
```

```python
grader.check("task_02")
```

### Task 03 📍

Define a function `correlation` that computes the correlation between 2 arrays of data in original units.

**Hint:** Consider using the `standard_units` function you have defined previously.

_Points:_ 2

```python
def correlation(x, y):
    ...
```

```python
grader.check("task_03")
```

### Task 04 📍

Which of the following are true about the correlation coefficient?

1. The correlation coefficient measures the strength of a linear relationship.
2. When looking at the existing data, a correlation coefficient of 1.0 means an increase in one variable always means an increase in the other variable.
3. The correlation coefficient is the slope of the regression line in standard units.
4. The correlation coefficient stays the same if we swap our x-axis and y-axis.
5. If we add 5 to our original data, our correlation coefficient will increase by 5.

Assign `r_array` to an array of your selections.

_Points:_ 2

```python
r_array = make_array(...)
```

```python
grader.check("task_04")
```

### Task 05 📍

Define a function `slope` that computes the slope of our line of best fit (to predict y given x), given two arrays of data in original units. Assume we want to create a line of best fit in original units.

**Hint:** Consider using the `correlation` function you have defined previously.

_Points:_ 2

```python
def slope(x, y):
    r = ...
    ...
```

```python
grader.check("task_05")
```

### Task 06 📍

Which of the following are true statements about the slope of our line of best fit? Assume $x$ refers to the value of one variable that we use to predict the value of y.

1. In original units, the slope has the unit: unit of $x$ / unit of $y$.
2. In standard units, the slope is unitless.
3. In original units, the slope is unchanged by swapping $x$ and $y$.
4. In standard units, a slope of 1 means our data is perfectly linearly correlated.
5. In original units and standard units, the slope always has the same positive or negative sign.
 
Assign `slope_array` to an array of your selections.

_Points:_ 2

```python
slope_array = make_array(...)
```

```python
grader.check("task_06")
```

### Task 07 📍

Define a function `intercept` that computes the intercept of our line of best fit (to predict `y` given `x`), given 2 arrays of data in original units. Assume we want to create a line of best fit in original units.

**Hint:** Consider using slope function you have defined previously.

_Points:_ 2

```python
def intercept(x, y):
    ...
```

```python
grader.check("task_07")
```

### Task 08 📍

Which of the following are true about the intercept of our line of best fit? Assume $x$ refers to the value of one variable that we use to predict the value of $y$.

1. In original units, the intercept has the same unit as the $y$ values.
2. In original units, the intercept has the same unit as the $x$ values.
3. In original units, the slope and intercept have the same unit.
4. In standard units, the intercept for the regression line is 0.
5. In original units and standard units, the intercept always has the same numerical value.

Assign `intercept_array` to an array of your selections.

_Points:_ 2

```python
intercept_array = make_array(...)
```

```python
grader.check("task_08")
```

### Task 09 📍

Define a function `predict` that takes in a table and 2 column names, and returns an array of predictions. Keep the following in mind when you make this function:
* The predictions should be created using a fitted regression line.
* We are predicting `y_col` from `x_col`, both in original units.

**Hint 1:** Consider using the `slope` and `intercept` functions you have defined previously.

**Hint 2:** Re-reading [Section 15.2](https://inferentialthinking.com/chapters/15/2/Regression_Line.html) might be helpful here.

**Note:** The public tests are quite comprehensive for this question, so passing them means that your function most likely works correctly.

_Points:_ 3

```python
def predict(tbl, x_col, y_col):
    x = ...
    y = ...
    ...
```

```python
grader.check("task_09")
```

## FIFA Predictions

The following data was [scraped](https://en.wikipedia.org/wiki/Data_scraping) from [sofifa.com](https://sofifa.com/), a website dedicated to collecting information from FIFA video games. The dataset consists of all players in [FIFA 22](https://en.wikipedia.org/wiki/FIFA_22) and their corresponding attributes. We have truncated the dataset to a limited number of rows (100) to ease with our visualizations and analysis. 

**Since we're learning about linear regression, we will look specifically for a linear association between various player attributes. To help with understanding where the line of best fit generated in linear regression comes from please do not use the `.fit_line` argument in `.scatter` at any point on question 2 unless the code was provided for you.**

Run the following cell to load the data and reduce the original data to a more manageable amount of information. (There are 110 columns in the original dataset!) 

```python
fifa = Table.read_table('fifa22.csv')
fifa = fifa.select("short_name", "overall", "value_eur", 
                   "wage_eur", "age", "pace", "shooting", 
                   "passing", "attacking_finishing")
fifa.show(5)
```

### Task 10 📍🔎

<!-- BEGIN QUESTION -->

Before jumping into any statistical techniques, it's important to see what the data looks like, because data visualizations allow us to uncover patterns in our data that would have otherwise been much more difficult to see.

Create a scatter plot with age on the x-axis (`"age"`), and the player's value in Euros (`"value_eur"`) on the y-axis.

_Points:_ 3

```python
...
```

<!-- END QUESTION -->

### Task 11 📍🔎

<!-- BEGIN QUESTION -->

Create a scatter plot with player age (`"age"`) along the x-axis and both real player value (`"value_eur"`) and predicted player value along the y-axis. 

* The predictions should be created using a fitted regression line.
* The color of the dots for the real player values should be different from the color for the predicted player values.

**Hint 1:** Consider using the `predict` function you have defined previously.

**Hint 2:** [Section 15.2](https://inferentialthinking.com/chapters/15/2/Regression_Line.html) and [Section 7.3](https://inferentialthinking.com/chapters/07/3/Overlaid_Graphs.html) have examples of creating such scatter plots.

_Points:_ 3

```python
predictions = ...
fifa_with_predictions = ...
...
```

<!-- END QUESTION -->

### Task 12 📍🔎

<!-- BEGIN QUESTION -->

Looking at the scatter plot you produced above, is linear regression a good model to use? If so, what features or characteristics make this model reasonable? If not, what features or characteristics make it unreasonable?

_Points:_ 3

_Type your answer here, replacing this text._

<!-- END QUESTION -->

### Task 13 📍

Define a function `rmse` that takes in two arguments: a slope and an intercept for a potential regression line. The function should return the root mean squared error between the values predicted by a regression line with the given slope and intercept and the actual outcomes.

**Note:** Assume we are still predicting `"value_eur"` from `"age"` in original units from the `fifa` table.

_Points:_ 2

```python
def rmse(slope, intercept):
    predictions = ...
    errors = ...
    ...
```

```python
grader.check("task_13")
```

### Task 14 📍🔎

<!-- BEGIN QUESTION -->

1. Use the `rmse` function you defined along with `minimize` to find the least-squares regression parameters predicting player value from player age. Review the [example of using the `minimize` function from the textbook](https://inferentialthinking.com/chapters/15/3/Method_of_Least_Squares.html#numerical-optimization).
2. Then set `lsq_slope` and `lsq_intercept` to be the least-squares regression line slope and intercept, respectively.
3. Finally, create a scatter plot like you did in Task 11 with player age (`"age"`) along the x-axis and both real player value (`"value_eur"`) and predicted player value along the y-axis. Be sure to use your least-squares regression line to compute the predicted values. The color of the dots for the real player values should be different from the color for the predicted player values.

**Note:** Your solution should **not** make any calls to the slope or intercept functions defined earlier. You are demonstrating a different method for creating the line in this task.

**Hint:** Your call to `minimize` will return an array of argument values that minimize the return value of the function passed to minimize.

_Points:_ 3

```python
minimized_parameters = ...
lsq_slope = ...
lsq_intercept = ...

# This just prints your slope and intercept
print("Slope: {:g} | Intercept: {:g}".format(lsq_slope, lsq_intercept))

lsq_predictions = ...
fifa_with_lsq_predictions = ...
...
```

<!-- END QUESTION -->

Run this cell to generate a scatter plot for the next part.

```python
no_goalies = fifa.where("shooting", are.above(0))
no_goalies.scatter('shooting', 'attacking_finishing', fit_line=True)
```

### Task 15 📍

Above is a scatter plot showing the relationship between a player's shooting ability (`"shooting"`) and their scoring ability (`"attacking_finishing"`). This visual excludes information on goalies. Goalies do not have shooting information in this dataset, but they have scoring ability data, so we've removed the goalies from the `fifa` table when producing the above image.

There is clearly a strong positive correlation between the two variables, and we'd like to predict a player's scoring ability from their shooting ability. 

Which of the following statements are true, assuming linear regression is a reasonable model?

**Hint:** Re-reading [Section 15.2](https://inferentialthinking.com/chapters/15/2/Regression_Line.html) might be helpful here.

1. For a majority of players with a shooting attribute above 80 our model predicts they have a better scoring ability than shooting ability.
2. A randomly selected player's predicted scoring ability in standard units will always be less than their shooting ability in standard units.
3. If we select a player who's shooting ability is 1.0 in standard units, their scoring ability, on average, will be less than 1.0 in standard units.
4. Goalies have `attacking_finishing` scores in our dataset but do not have shooting scores. We can still use our model to predict their attacking_finishing scores.

Assign `scoring_array` to an array of your selections.

_Points:_ 2

```python
scoring_array = make_array(...)
```

```python
grader.check("task_15")
```

## Submit your Homework to Canvas

Once you have finished working on the homework tasks, prepare to submit your work in Canvas by completing the following steps.

1. In the related Canvas Assignment page, check the rubric to know how you will be scored for this assignment.
2. Double-check that you have run the code cell near the end of the notebook that contains the command `"grader.check_all()"`. This command will run all of the run tests on all your responses to the auto-graded tasks marked with 📍.
3. Double-check your responses to the manually graded tasks marked with 📍🔎.
3. Select the menu item "File" and "Save Notebook" in the notebook's Toolbar to save your work and create a specific checkpoint in the notebook's work history.
4. Select the menu items "File", "Download" in the notebook's Toolbar to download the notebook (.ipynb) file. 
5. In the related Canvas Assignment page, click Start Assignment or New Attempt to upload the downloaded .ipynb file.

**Keep in mind that the auto-grader does not always check for correctness. Sometimes it just checks for the format of your answer, so passing the auto-grader for a task does not mean you got the answer correct for that task.**

---

To double-check your work, the cell below will rerun all of the autograder tests.

```python
grader.check_all()
```

Tests for task_01 (points: 2)
```
>>> standard_units(make_array(1, 2, 3, 4, 5))
array([-1.41421356, -0.70710678,  0.        ,  0.70710678,  1.41421356])
```
failure_message: ❌ standard_units doesn't work for the test data we provided.
success_message: ✅ standard_units works for the test data we provided.
```
>>> np.isclose(np.mean(standard_units([1, 2, 3, 4, 5])), 0)
True
```
failure_message: ❌ The mean of the standardized test data should be 0.
success_message: ✅ The mean of the standardized test data is 0.
```
>>> np.isclose(np.std(standard_units([1, 2, 3, 4, 5])), 0.9999999999999999)
True
```
failure_message: ❌ The standard deviation of the standardized test data should be 1.
success_message: ✅ The standard deviation of the standardized test data is 1.


Tests for task_02 (points: 2)
```
>>> set(standard_array).issubset({1, 2, 3, 4, 5})
True
```
failure_message: ❌ standard_array should contain some of the integers 1, 2, 3, 4, or 5.
success_message: ✅ standard_array contains some of the integers 1, 2, 3, 4, or 5.


Tests for task_03 (points: 2)
```
>>> np.isclose(correlation([1, 2, 3], [4, 5, 6]), 0.9999999999999999)
True
```
failure_message: ❌ correlation should show that linear data has a coefficient of 1.
success_message: ✅ correlation shows that linear data has a coefficient of 1.


Tests for task_04 (points: 2)
```
>>> set(r_array).issubset({1, 2, 3, 4, 5})
True
```
failure_message: ❌ r_array should contain some of the integers 1, 2, 3, 4, or 5.
success_message: ✅ r_array contains some of the integers 1, 2, 3, 4, or 5.


Tests for task_05 (points: 2)
```
>>> test_data = make_array(1, 2, 3, 4)
>>> np.isclose(slope(test_data, 2 * test_data), 2)
True
```
failure_message: ❌ slope did not produce the correct result for our test data set.
success_message: ✅ slope produced the correct result for our test data set.


Tests for task_06 (points: 2)
```
>>> set(slope_array).issubset({1, 2, 3, 4, 5})
True
```
failure_message: ❌ slope_array should contain some of the integers 1, 2, 3, 4, or 5.
success_message: ✅ slope_array contains some of the integers 1, 2, 3, 4, or 5.


Tests for task_07 (points: 2)
```
>>> test_data = make_array(1, 2, 3, 4)
>>> np.isclose(intercept(test_data, 2 * test_data), 0)
True
```
failure_message: ❌ intercept did not produce the correct result for our test data set.
success_message: ✅ intercept produced the correct result for our test data set.


Tests for task_08 (points: 2)
```
>>> set(intercept_array).issubset({1, 2, 3, 4, 5})
True
```
failure_message: ❌ intercept_array should contain some of the integers 1, 2, 3, 4, or 5.
success_message: ✅ intercept_array contains some of the integers 1, 2, 3, 4, or 5.


Tests for task_09 (points: 3)
```
>>> test_data = make_array(1, 2, 3, 4)
>>> test_tbl = Table().with_columns('x_col', test_data, 'y_col', 2 * test_data)
>>> test_prediction = predict(test_tbl, 'x_col', 'y_col')
>>> isinstance(test_prediction, np.ndarray)
True
```
failure_message: ❌ predict should return an array.
success_message: ✅ predict returns an array.
```
>>> test_data = make_array(1, 2, 3, 4)
>>> test_tbl = Table().with_columns('x_col', test_data, 'y_col', 2 * test_data)
>>> predict(test_tbl, 'x_col', 'y_col')
array([ 2.,  4.,  6.,  8.])
```
failure_message: ❌ predict seems to work correctly on the test data we used.
success_message: ✅ predict seems to work correctly on the test data we used.
```
>>> test_data = make_array(-2, -1, 1, 2)
>>> test_tbl = Table().with_columns('x_col', test_data, 'y_col', test_data)
>>> predict(test_tbl, 'x_col', 'y_col')
array([-2., -1.,  1.,  2.])
```
failure_message: ❌ predict seems to work correctly on the test data we used.
success_message: ✅ predict seems to work correctly on the test data we used.
```
>>> test_data = make_array(-2, -1, 1, 2)
>>> test_tbl = Table().with_columns('x_col', test_data, 'y_col', -1 * test_data)
>>> predict(test_tbl, 'x_col', 'y_col')
array([ 2.,  1., -1., -2.])
```
failure_message: ❌ predict seems to work correctly on the test data we used.
success_message: ✅ predict seems to work correctly on the test data we used.


Tests for task_13 (points: 2)
```
>>> max_value_eur = np.max(fifa.column('value_eur'))
>>> min_value_eur = np.min(fifa.column('value_eur'))
>>> max_age = np.max(fifa.column('age'))
>>> min_age = np.min(fifa.column('age'))
>>> test_slope = (max_value_eur - min_value_eur) / (max_age - min_age)
>>> test_intercept = 1.2 * 100000000.0
>>> test_rmse = rmse(test_slope, test_intercept)
>>> isinstance(test_rmse, (float, np.float_))
True
```
failure_message: ❌ rmse should return a float.
success_message: ✅ rmse returns a float.
```
>>> max_value_eur = np.max(fifa.column('value_eur'))
>>> min_value_eur = np.min(fifa.column('value_eur'))
>>> max_age = np.max(fifa.column('age'))
>>> min_age = np.min(fifa.column('age'))
>>> test_slope = (max_value_eur - min_value_eur) / (max_age - min_age)
>>> test_intercept = 1.2 * 100000000.0
>>> np.isclose(rmse(test_slope, test_intercept), 346627863.32058465)
True
```
failure_message: ❌ rmse does not seem to work on our test data.
success_message: ✅ rmse seems to work on our test data.


Tests for task_15 (points: 2)
```
>>> True
True
```
failure_message: ❌ 
success_message: ✅ 
