```python
# Initialize Otter
import otter
grader = otter.Notebook("hw04.ipynb")
```

<img style="display: block; margin-left: auto; margin-right: auto" src="./ccsf-logo.png" width="250rem;" alt="The CCSF black and white logo">

# Homework 4: Data Analysis

## References

* [Chapter 8 - Functions and Tables](https://inferentialthinking.com/chapters/08/Functions_and_Tables.html#) 
* [`datascience` Documentation](https://datascience.readthedocs.io/)
* [Python Quick Reference](https://ccsf-math-108.github.io/materials-fa23/resources/quick_reference.html)

## Assignment Reminders

- Make sure to run the code cell at the top of this notebook that starts with `# Initialize Otter` to load the auto-grader.
- For all tasks indicated with a 🔎 that you must write explanations and sentences for, provide your answer in the designated space.
- Throughout this assignment and all future ones, please be sure to not re-assign variables throughout the notebook! _For example, if you use `max_temperature` in your answer to one question, do not reassign it later on. Otherwise, you will fail tests that you thought you were passing previously!_
- We encourage you to discuss this assignment with others but make sure to write and submit your own code. Refer to the syllabus to learn more about how to learn cooperatively.
- Unless you are asked otherwise, use the non-interactive visualizations when asked to produce a visualization for a task.
- View the related <a href="https://ccsf.instructure.com" target="_blank">Canvas</a> Assignment page for additional details.

Run the following code cell to import the tools for this assignment.

```python
from datascience import *
import numpy as np
%matplotlib inline
import matplotlib.pyplot as plt
plt.style.use('fivethirtyeight')
```

## 🌉 San Francisco City and County Employee Salaries

In this homework assignment, you will combine several of the skills and concepts you've learned so far to analyze [Employee Compensation](https://data.sfgov.org/City-Management-and-Ethics/Employee-Compensation/88g8-5mnd) data provided by the SF Controller's Office. We have filtered it to retain just the relevant columns and restricted the data to the fiscal year 2022-2023 (July 1, 2022 - June 30, 2023). 

The data stored on your account was pulled from [Employee Compensation page](https://data.sfgov.org/City-Management-and-Ethics/Employee-Compensation/88g8-5mnd) at `DataSF.gov` using the API endpoint: 

`'https://data.sfgov.org/resource/88g8-5mnd.csv?$limit=1000000&year_type=Fiscal&year=2023'`

Run the following cell to load that data into a table called `sf`.

```python
# API endpoint: https://data.sfgov.org/resource/88g8-5mnd.csv?$limit=1000000&year_type=Fiscal&year=2023
sf = Table.read_table('sf_fy22_23.csv')
sf
```

### 📝 Data Notes

* Each line represents an employee's job information such as job family, salary, and benefits.
* The San Francisco Controller's Office  provides [a PDF explaining what each variable means](https://data.sfgov.org/api/views/88g8-5mnd/files/OMBVvreoXRjXG6oP4Ts4497dNxt14XlBqB6uIL6cq-o?download=True&filename=N:\EIS\DataCoordination\Metadata%20Spring%20Cleaning\CON_DataDictionary_Employee-Compensation.pdf).
There are a few interesting details in the data set that are explained in the [Employee Compensation FAQ](https://support.datasf.org/help/employee-compensation-faq)
    * Employees have been given an employee identifier for this data and is not connected to their actual employee identification.
    * An employee can appear on multiple lines
    * Rows may have zero dollar amounts
    * Adjustments to employee's pay can result in negative dollar amounts

### 🔔 Table Size Warning

This table has over 40 thousand rows in it. Just loading the data takes up more memory than most of our data sets so far. This means that you want to be careful with creating too many copies of the table since you have a limited amount of memory. We've set up the assignment, so you can complete the tasks with the provided amount of memory, but you may accidentally crash the kernel by running out of memory.

**If your kernel keeps crashing because of running out of memory:**
* Make sure you are only running this notebook and all your other kernels are shut down.
* Talk with a staff member about things you might be doing that utilize a lot of resources.

### Task 01 📍

The [Employee Compensation FAQ](https://support.datasf.org/help/employee-compensation-faq) states:
> One employee can appear on multiple lines per year if they changed positions and/or departments mid-year, which happens with some frequency.

This may or may not be true for this particular fiscal year. There are 41,295 rows in the `sf` table. Determine how many unique employee identifiers there are in the `sf` table to see if there are any employees listed more than once. Assign `num_employee_ids` to the number (`int`) of unique employee identifiers in the `sf` table.

*Hint: Consider using the `group` table method.*

_Points:_ 2

```python
num_employee_ids = ...
num_employee_ids
```

```python
grader.check("task_01")
```

### Task 02 📍🔎

<!-- BEGIN QUESTION -->

There are a few departments that have a large number of employees. For this task, produce a bar chart that shows the five largest departments in the `sf` table ranked by employee count.

_Points:_ 2

```python
...

# Leave this to provide a title for your visualization
plt.title('Top 5 Departments By Employee Count')
plt.show()
```

<!-- END QUESTION -->

### Task 03 📍

Some employees for the city and county make a lot of money. Identify the job title that makes the most in terms of total compensation (the sum of salary and benefits).
1. Assign the largest total compensation (a `float` value) to `max_compensation`.
2. Assign the job title (a `str` value) associated with the highest total compensation to the name `max_compensation_position`.

_Points:_ 4

```python
max_compensation = ...
max_compensation_position = ...
print(f"The largest total compensation is ${max_compensation:,}")
print(f"The position associated with this compensation is: {max_compensation_position}")
```

```python
grader.check("task_03")
```

### Task 04 📍🔎

<!-- BEGIN QUESTION -->

Next, summarize the distribution of total compensation values so you can compare that large salary with the salaries of the rest of the employees. Specifically, create a histogram showing the distribution of total compensation.
* Use the `unit = '$'` argument
* Use the default bins.

_Points:_ 2

```python
...

# Leave this to provide a title for your visualization
plt.title('Total Compensation')
plt.show()
```

<!-- END QUESTION -->

### Task 05 📍🔎

<!-- BEGIN QUESTION -->

In the provided Markdown cell, explain why there seems to be no information visualized on the right side of the histogram created in Task 04. Make sure you connect your response to the `max_compensation` value you calculated.

_Points:_ 2

_Type your answer here, replacing this text._

<!-- END QUESTION -->

### Task 06 📍🔎

<!-- BEGIN QUESTION -->

The area of a bar in a histogram reflects the percentage of the data represented in that particular bin. To have you think more deeply about this, we want you to create a strange version of the histogram you just made.

1. Create an array called `equal_split_bins` that can be used to form the bins for a histogram of total compensation that splits the data into two equal bins.
    * _Hint: The median (`np.median`) of a collection of numbers can be used to split the numbers into two equal halves._
    * _Hint: Don't forget that you named the largest total compensation as `max_compensation` in a previous task._
2. Create a histogram for the total compensation values by using `equal_split_bins` such that your histogram only shows 2 bins (bars).
    * _Hint: Use the `bins` argument for `hist` to incorporate the `equal_split_bins` array._

As with Task 02, we are just scoring your visualization, so you are welcome to alter the template/steps.

_Points:_ 2

```python
# Step 1
min_compensation = ...
median_compensation = ...
equal_split_bins = ...

# Step 2
...

# Leave this to provide a title for your visualization
plt.title('Total Compensation')
plt.show()
```

<!-- END QUESTION -->

Now, that you have a basic feel for the distribution of total compensation values. Summarize the total compensations based on department. Salaries may vary within a department depending on several factors such as job, experience, etc. There might be extremely high or low salaries within a department, so it is best to use a summary value like the median which is less influenced by extreme values. 

### Task 07 📍

Your goal for this task is to create a table titled `top_5_departments` featuring two columns: `department` listing department names and `total_compensation median` displaying the median total compensation for each department. Once done, the final table should showcase the 5 departments in the data set with the highest median total compensation values.

_Points:_ 5

```python
top_5_departments = ...
top_5_departments
```

```python
grader.check("task_07")
```

### Task 08 📍

Set `department_ranges` to a table containing departments as the rows, and the organization groups as the columns. The values in the rows should correspond to a total compensation range, where the range is defined as the **difference between the highest total compensation and the lowest total compensation in the department for that organization group**.

Keep in mind the following while working on this:

* First you'll need to define a new function `compensation_range` which takes in an array of compensations and returns the range of compensations in that array.
* What table function allows you to specify the rows and columns of a new table?


_Points:_ 4

```python
# Define compensation_range first
...

department_ranges = ...
department_ranges
```

```python
grader.check("task_08")
```

### Task 09 📍🔎

<!-- BEGIN QUESTION -->

Provide at least **two** different explanations as to why some of the row values are `0` in the `department_ranges` table from the previous question.


_Points:_ 2

_Type your answer here, replacing this text._

<!-- END QUESTION -->

Which department has the longest job title? There are a few ways you might answer this, but we are going to guide you in one way that gives you practice making and using functions with the `group` method. 

Before you go on to the next task, run the following code cell to define a function that takes a list or array of strings and returns an array of the lengths of those strings. You don't have to understand how the code works, but this might give you an idea of where error messages come from. You should try it out with a few lists/arrays of strings to see what it is doing.

```python
def string_lengths(strings):
    '''
    Returns an array that contains the length of each string in the array or list called strings.

    Example
    --------
    >>> string_lengths(['a', 'an', 'the'])
    array([1, 2, 3])
    '''
    
    if not isinstance(strings, (list, np.ndarray)):
        raise ValueError("Input must be a list or array")

    if len(strings) == 0:
        return np.array([0])
    
    for s in strings:
        if not isinstance(s, str):
            raise ValueError("Input must contain only strings")
        
    return np.array([len(s) for s in strings])
```

```python
# Try calling the function
an_array_of_strings = make_array('City', 'College', 'of', 'San Francisco')
string_lengths(an_array_of_strings)
```

Notice that it provides the length of the 4 strings that were provided in the input array. Remember that `len` does count the blank space in `'San Francisco'`.

### Task 10 📍

Define a function called `max_job_name` that returns the length of the longest job name from an array or list of job names (strings).

* **Hint:** Use the `string_lengths` function above. 
* **Hint:** `max_job_name(['Electrician', 'Analyst', 'Social Worker'])` should return `13` because `'Social Worker'` is the longest string with 13 characters (including the space).
* **Hint:** You can assume the list of jobs is not empty because every department in the `sf` table has at least 1 job. 

_Points:_ 3

```python
...

# Try out the function
max_job_name(['Electrician', 'Analyst', 'Social Worker'])
```

```python
grader.check("task_10")
```

### Task 11 📍

What are the lengths of the longest job names for each department?

Create a table called `departments` that shows a row for each unique department name in the `sf` table and the length of the longest job name within that department. A few things to consider when making your final table:

* Sort the table so that the department(s) with the longest job name are at the top of the table.
* The columns should be `'department'` and `'max_job_name_length'`.

**Hint:** Try out the `max_job_name` function you just made!

_Points:_ 2

```python
departments =
departments
```

```python
grader.check("task_11")
```

---

California has laws in place to help govern how much an employee should be paid for overtime work. Visit the [Labor Commissioner's Office FAQ about Overtime](https://www.dir.ca.gov/dlse/faq_overtime.htm) if you are curious about these laws. There could be several reasons for why an employee works overtime. Some are healthy reasons and some are not. The ratio of overtime compensation to total compensation can provide a signal for health of departments and their employees. According to [Indeed.com](https://www.indeed.com/career-advice/career-development/working-overtime), here are some potential disadvantages of working extra hours:

> **Focus loss**: You will likely want to take breaks while working overtime, and you may lose focus and productivity naturally as your working hours increase. 
> 
> **Safety and health risk**: Working longer hours also can be dangerous, depending on the job. Working overtime regularly can also disrupt your work-life balance, lead to burnout or create health risks, such as sitting at a computer for long periods. Due to these risks, more companies are limiting the number of hours worked in certain positions, such as truck drivers. 
> 
> **Less work-life balance**: There are only 24 hours in the day, and working overtime reduces the time for a good work-life balance. More work hours mean fewer hours for family, relaxation and sleep.

On the other hand, here are some potential advantages for working extra hours:

> **Extra money**: Working overtime means you earn extra money for extra hours, usually 1.5 to two times your hourly rate. Overtime pay could allow you to bulk up your salary and have more money for living expenses, vacations or savings. 
>
>**Career boost**: Working overtime may also help advance your career—and potential earnings—by proving you’re reliable and motivated in your role. Volunteering to work extra hours shows your current employer that you’re dedicated to doing a good job. 
>
>**Gain experience**: Overtime is an opportunity to gain job experience and polish your skills. Typically, the more you work at something, the better you become.

One way to analyze overtime work for San Francisco employees is to consider the overtime pay values in the data set, but these values are relative to the amount of money the employee makes. To account for this you could consider overtime pay for an employee as a quantity relative to that employee's salary. This would allow you to compare these ratios for all the employees. Specifically, use the measurement called the overtime ratio defined as:

$$\text{Overtime Ratio} = \frac{\text{Overtime Pay}}{\text{Salary}}$$

There are many special cases when it comes to overtime pay that might analyzing this ratio in general challenging. So, to keep things more simple, narrow your focus to employees to at least a certain salary. Let that base salary reflect working 40 hrs per week for 52 weeks at the minimum wage of the time, 16.99 dollars per hour.

### Task 12 📍

For this task, create a table called `sf_above_35k` that contains the same information in the `sf` table and overtime ratios for employees making above $35,000. Make sure that `sf_above_35k` has:
 * Only employees from the `sf` table that have salaries above 35,000 dollars AND
 * The same columns as the `sf` table with an additional column `overtime_ratio` containing the overtime ratios from the above formula.

_Points:_ 5

```python
overtime_ratios = ...

sf_above_35k = ...
sf_above_35k
```

```python
grader.check("task_12")
```

### Task 13 📍🔎

<!-- BEGIN QUESTION -->

Is overtime collected similarly over the various salaries in the `sf_above_35k` table? See if there is any relationship between salary and the overtime ratios. For this task, create a visualization showing the relationship between salary and overtime ratio values for employees who make above $35,000.

_Points:_ 2

```python
...

# Leave this to provide a title for your visualization
plt.title('Overtime Ratio vs. Salary')
plt.show()
```

<!-- END QUESTION -->

### Task 14 📍🔎

<!-- BEGIN QUESTION -->

In the provided Markdown cell, summarize the relationship you see based on the visualization and provide at least one idea for why you are seeing the pattern you are seeing.

_Points:_ 2

_Type your answer here, replacing this text._

<!-- END QUESTION -->

There are many things you can explore in this data set, but that is enough for now!

## Submit your Homework to Canvas

Once you have finished working on the homework tasks, prepare to submit your work in Canvas by completing the following steps.

1. In the related Canvas Assignment page, check the rubric to know how you will be scored for this assignment.
2. Double-check that you have run the code cell near the end of the notebook that contains the command `"grader.check_all()"`. This command will run all of the run tests on all your responses to the auto-graded tasks marked with 📍.
3. Double-check your responses to the manually graded tasks marked with 📍🔎.
3. Select the menu item "File" and "Save Notebook" in the notebook's Toolbar to save your work and create a specific checkpoint in the notebook's work history.
4. Select the menu items "File", "Download" in the notebook's Toolbar to download the notebook (.ipynb) file. 
5. In the related Canvas Assignment page, click Start Assignment or New Attempt to upload the downloaded .ipynb file.

**Keep in mind that the autograder does not always check for correctness. Sometimes it just checks for the format of your answer, so passing the autograder for a question does not mean you got the answer correct for that question.**

---

To double-check your work, the cell below will rerun all of the autograder tests.

```python
grader.check_all()
```

Tests for task_01 (points: 2)
```
>>> isinstance(num_employee_ids, (int, np.integer))
True
```
failure_message: ❌ num_employee_ids should represent an integer.
success_message: ✅ num_employee_ids represents an integer.


Tests for task_03 (points: 4)
```
>>> isinstance(max_compensation, (float, np.float64, np.float32))
True
```
failure_message: ❌ max_compensation should represent a float.
success_message: ✅ max_compensation represents a float.
```
>>> isinstance(max_compensation_position, str)
True
```
failure_message: ❌ max_compensation_position should represent a string.
success_message: ✅ max_compensation_position represents a string.


Tests for task_07 (points: 5)
```
>>> isinstance(top_5_departments, Table)
True
```
failure_message: ❌ top_5_departments should be a Table.
success_message: ✅ top_5_departments is a Table.
```
>>> top_5_departments.num_rows == 5 and top_5_departments.num_columns == 2
True
```
failure_message: ❌ top_5_departments should have 5 rows and 2 columns.
success_message: ✅ top_5_departments has the correct number of rows and columns


Tests for task_08 (points: 4)
```
>>> compensation_range(make_array(5, 1, 20, 1000)) == 999
True
```
failure_message: ❌ Your function compensation_range doesn't seem to be working correctly.
success_message: ✅ Your function compensation_range seems to be working correctly.
```
>>> set(['Department', 'Culture & Recreation', 'General Administration & Finance', 'Human Welfare & Neighborhood Development', 'Public Protection', 'Public Works, Transportation & Commerce']) == set(department_ranges.labels)
False
```
failure_message: ❌ You don't seem to have the correct column labels in department_ranges.
success_message: ✅ You have the correct column labels in department_ranges.
```
>>> sum(department_ranges.column(1)) == 589454.63
False
```
failure_message: ❌ The values in one of the columns of department_ranges don't seem correct.
success_message: ✅ The values in one of the columns of department_ranges seem correct.


Tests for task_10 (points: 3)
```
>>> callable(max_job_name)
True
```
failure_message: ❌ max_job_name is not callable.
success_message: ✅ max_job_name is callable.
```
>>> max_job_name(['Electrician', 'Analyst', 'Social Worker']) == 13
True
```
failure_message: ❌ max_job_name doesn't seem to be working.
success_message: ✅ max_job_name seems to be working.


Tests for task_11 (points: 2)
```
>>> def check_old_new():
...     try:
...         return isinstance(departments, Table) or isinstance(department, str)
...     except NameError:
...         if 'departments' not in globals() and ('department' in globals() and isinstance(department, str)):
...             return True
...         else:
...             return False
...     else:
...         return False
>>> check_old_new()
True
```
failure_message: ❌ departments should a table.
success_message: ✅ departments is a table (or department is a string).


Tests for task_12 (points: 5)
```
>>> isinstance(sf_above_35k, Table)
True
```
failure_message: ❌ sf_above_35k should be a Table.
success_message: ✅ sf_above_35k is a Table.
```
>>> sf_above_35k.num_columns == sf.num_columns + 1
True
```
failure_message: ❌ sf_above_35k doesn't have the correct number of columns.
success_message: ✅ sf_above_35k has the correct number of columns.
