```python
# Initialize Otter
import otter
grader = otter.Notebook("hw03.ipynb")
```

<img style="display: block; margin-left: auto; margin-right: auto" src="./ccsf-logo.png" width="250rem;" alt="The CCSF black and white logo">

# Homework 03: Visualizations

**Recommended Reading**: 
* [Tables](https://inferentialthinking.com/chapters/06/Tables.html)
* [Visualization](https://inferentialthinking.com/chapters/07/Visualization.html)

## Assignment Reminders

- Make sure to run the code cell at the top of this notebook that starts with `# Initialize Otter` to load the auto-grader.
- For all tasks indicated with a 🔎 that you must write explanations and sentences for, provide your answer in the designated space.
- Throughout this assignment and all future ones, please be sure to not re-assign variables throughout the notebook! _For example, if you use `max_temperature` in your answer to one question, do not reassign it later on. Otherwise, you will fail tests that you thought you were passing previously!_
- We encourage you to discuss this assignment with others, but make sure to write and submit your own code. Refer to the syllabus to learn more about how to learn cooperatively.

*View the related <a href="https://ccsf.instructure.com" target="_blank">Canvas</a> Assignment page for additional details.*

Run the following code cell to import the tools for this assignment.

```python
import numpy as np
from datascience import *
import matplotlib
%matplotlib inline
import matplotlib.pyplot as plt
plt.style.use('fivethirtyeight')
```

## Unemployment


The [Federal Reserve Bank of St. Louis](https://fred.stlouisfed.org/categories/33509) publishes data about jobs in the US.  Below, we've loaded data on unemployment in the United States. There are many ways of defining unemployment, and our dataset includes two notions of the unemployment rate:

1. Among people who are able to work and are looking for a full-time job, the percentage who can't find a job.  This is called the Non-Employment Index, or NEI.
2. Among people who are able to work and are looking for a full-time job, the percentage who can't find any job *or* are only working at a part-time job.  The latter group is called "Part-Time for Economic Reasons", so the acronym for this index is NEI+PTER.  (Economists are great at marketing.)

The data that you will use in this assignment contains [quarterly average NEI percentages from 1994 until 2024](https://fred.stlouisfed.org/series/NEIM156SFRBRIC) and [quarterly average NEI+PTER percentages from 1994 until 2024](https://fred.stlouisfed.org/series/NEIPTERM156SFRBRIC)

*In this assignment, you will do some preliminary analysis of these unemployment metrics. In a future assignment, you will work on creating visualizations and summarizing your analysis.*


### Task 01 📍

The data are in a CSV file called `unemployment.csv`.  Load the data in that file into a `Table` and assign the table to the name `unemployment`.


_Points:_ 1

```python
unemployment = ...
unemployment
```

```python
grader.check("task_01")
```

### Task 02 📍

Dates can be tricky to work with. When you loaded the data into the table `unemployment`, the dates were read as strings. Run the following cell to see this is true.

```python
type(unemployment.column('DATE').item(0))
```

Strings do not always make the best sense to use for dates since we organize dates by months, years, etc. and this is different than the alphabetical order we typically use for strings. A common date type for dates is [`datetime`](https://docs.python.org/3/library/datetime.html). In this class, you are not responsible for knowing how to work directly with the `datetime` format. Instead, we will set up the code to help you with that date type.

**Run the following code cell to convert all the dates in the `'DATE'` column to `datetime` format.**

_Points:_ 0

```python
# Just run this cell
import datetime
from datetime import datetime, date

def update_date_format(date_string):
    date_object = datetime.strptime(date_string, '%Y-%m-%d').date()
    return date_object

if not isinstance(unemployment.column('DATE').item(0),  date):
    unemployment = unemployment.with_column(
        'DATE', unemployment.apply(update_date_format, 'DATE'))
    print("The date type for the values in the 'DATE' column have been updated.\n")
else:
    print("The date type is correct for the values in the 'DATE' column.\n")
    
display(unemployment)
```

```python
grader.check("task_02")
```

### Task 03 📍

For this task:
1. Sort the data in descending order by NEI, naming the sorted table `by_nei`.
2. Create another table called `by_nei_pter` that's sorted in descending order by NEI+PTER instead.


_Points:_ 2

```python
by_nei = ...
by_nei_pter = ...
```

```python
grader.check("task_03")
```

### Task 04 📍

For this task:
1. Use `take` to make a table containing the data for the 10 quarters when NEI was greatest.
2. Assign that table to the name `greatest_nei`. Keep in mind that `greatest_nei` should be sorted in descending order of `NEI`.

**Note**: Each row of `unemployment` represents a quarter.


_Points:_ 2

```python
greatest_nei = ...
greatest_nei
```

```python
grader.check("task_04")
```

### Task 05 📍

It's believed that many people became PTER (recall: "Part-Time for Economic Reasons") in the "Great Recession" of 2008-2009.  NEI+PTER is the percentage of people who are unemployed (and counted in the NEI) plus the percentage of people who are PTER.  

Using the original `unemployment` table, create an array containing the percentage of people who were PTER in each quarter.  (The first element of the array should correspond to the first row of `unemployment`, and so on.)

_Points:_ 2

```python
pter = ...
pter
```

```python
grader.check("task_05")
```

### Task 06 📍

Add `pter` as a column to `unemployment` (named `"PTER"`) and sort the resulting table by that column in descending order.  Assign this table to the name `unemployment` to update the original table.

**Notes**:
* You do not need to do this in one line of code.
* You are welcome to add and use extra variable names.
    * Make sure that you don't use a name that conflicts with the rest of the notebook.
    * Make sure that you use the name `unemployment` for the table since we will be checking that name with the auto-grader.

_Points:_ 2

```python
unemployment = ...
unemployment
```

```python
grader.check("task_06")
```

### Task 07 📍🔎

<!-- BEGIN QUESTION -->

Using the `unemployment` table, create an overlaid line plot showing two lines for NEI and PTER unemployment percentages over the dates. The dates should be on the horizontal axes and the percentages should be on the vertical axes.

**Note:** Our code below will format the graph for you so that the x-axis shows the year portion of the date for every 3 years starting from 1994.

_Points:_ 2

```python
...

# Leave the following code to improve the readability of the horizontal axis tick marks
start_date = min(unemployment.column("DATE"))
end_date = max(unemployment.column("DATE"))
years = [datetime(year, 1, 1) for year in range(start_date.year, end_date.year + 3, 3)]
plt.gca().set_xticks(years)
plt.gca().set_xticklabels([year.strftime('%Y') for year in years])
plt.xticks(rotation=45)
plt.title("Unemployment Percentages (1994 - 2024)")
plt.show()
```

<!-- END QUESTION -->

### Task 08 📍🔎

<!-- BEGIN QUESTION -->

As you saw early on in the course, sometimes the lines produced from data reveal a story in the data. For this task, we want you to review the graph for patterns, identify 3 major events in US history based on the graph, and hypothesize how the events and the patterns in the unemployment data are related.

In your response:
1. Mention 3 events that occurred between 1994 and 2024.
2. Express how those 3 events show up in the patterns you are observing.
3. Hypothesize how those events are associated with unemployment data.

_Points:_ 2

_Type your answer here, replacing this text._

<!-- END QUESTION -->

## Birth Rates


The CSV file `census.csv` contains census-based population estimates for each state on both July 1, 2022 and July 1, 2023. The data was taken from [the US Census 2020-2023 national totals data set](https://www2.census.gov/programs-surveys/popest/datasets/2020-2023/state/totals/NST-EST2023-ALLDATA.csv). 

Run the following code cell to load that data into a table called `pop`.

```python
pop = Table.read_table('census.csv')
pop
```

Here is a brief explanation of the column labels:

* `REGION`: Census Region code
* `NAME`: State name
* `'2022'`: 7/1/2022 resident total population estimate
* `'2023'`" 7/1/2023 resident total population estimate
* `'BIRTHS'`: Births in period 7/1/2022 to 6/30/2023			
* `'DEATHS'`: Deaths in period 7/1/2022 to 6/30/2023
* `'MIGRATION'`: Net migration in period 7/1/2022 to 6/30/2023
* `'OTHER'`: Residual for period 7/1/2022 to 6/30/2023

The last four columns describe the components of the estimated change in population during this time interval. 

**Note:** For all questions below, assume that the word "states" refers to all 52 rows including Puerto Rico & the District of Columbia.

### Task 09 📍

Assign `us_birth_rate` to the total US annual birth rate during this time interval. The annual birth rate for a year-long period is the total number of births in that period as a proportion of the population size at the start of the time period.

**Hint:** Which year corresponds to the start of the time period?


_Points:_ 2

```python
us_birth_rate = ...
us_birth_rate
```

```python
grader.check("task_09")
```

### Task 10 📍

Assign `movers` to the number of states for which the **absolute value** of the **annual rate of migration** was higher than 1%. The annual rate of migration for a year-long period is the net number of migrations (in and out) as a proportion of the population size at the start of the period. The `MIGRATION` column contains estimated annual net migration counts by state.


_Points:_ 2

```python
migration_rates = ...
movers = ...
movers
```

```python
grader.check("task_10")
```

### Task 11 📍

Assign `west_births` to the total number of births that occurred in region 4 (the Western US). 

**Hint:** Make sure you double check the type of the values in the region column, and appropriately filter (i.e. the types must match!).


_Points:_ 2

```python
west_births = ...
west_births
```

```python
grader.check("task_11")
```

### Task 12 📍

Assign `less_than_west_births` to the number of states that had a total population in 2023 that was smaller than the *total number of births in region 4 (the Western US)* during this time interval.


_Points:_ 2

```python
less_than_west_births = ...
less_than_west_births
```

```python
grader.check("task_12")
```

### Task 13 📍

In the next question, you will be creating a visualization to understand the relationship between birth and death rates for the states. The annual death rate for a year-long period in a state is the total number of deaths in that period for that state as a proportion of the population size at the start of the time period for that state.

What visualization is most appropriate to see if there is an association between birth and death rates during a given time interval among the states?

1. Line Graph
<br>
2. Scatter Plot
<br>
3. Bar Chart

Assign `visualization` below to the number corresponding to the correct visualization.


_Points:_ 2

```python
visualization = ...
```

```python
grader.check("task_13")
```

### Task 14 📍🔎

<!-- BEGIN QUESTION -->

In the code cell below, create a visualization that will help us determine if there is an association between birth rate and death rate during this time interval among the states. 

It may be helpful to create intermediate variables. In our template, we've introduced the names `birth_rates_2022` and `death_rates_2022` as suggestions. We will not test those names, so you do not need to use them. We will only score the visualization you produce.


_Points:_ 2

```python
birth_rates_2022 = ...
death_rates_2022 = ...

# Leave the following code to add a title to your graph
plt.title('2022 Death Rate vs. Birth Rate')
plt.show()
```

<!-- END QUESTION -->

### Task 15 📍

`True` or `False`: There is an association between birth rate and death rate during this time interval among the states. 

Assign `assoc` to `True` or `False` in the cell below. 


_Points:_ 2

```python
assoc = ...
```

```python
grader.check("task_15")
```

## Marginal Histograms


Consider the following scatter plot: 

<img src="scatter.png" alt="The Scatter plot" width=40%>

The axes of the plot represent values of two variables: $x$ and $y$. 

Suppose we have a table called `t` that has two columns in it:

- `x`: a column containing the x-values of the points in the scatter plot
- `y`: a column containing the y-values of the points in the scatter plot

Below, you are given two histograms, each of which corresponds to either column `x` or column `y`. 

### Histogram A

<img src="histogram_A.png" alt="Histogram A" width=40%>

### Histogram B

<img src="histogram_B.png" alt="Histogram B" width=40%>

### Task 16 📍

Suppose we run `t.hist('x')`. Which histogram does this code produce? Assign `histogram_column_x` to one of the following strings: `'A'`, `'B'`, or `'Neither'`.

_Points:_ 2

```python
histogram_column_x = ...
```

```python
grader.check("task_16")
```

### Task 17 📍

Suppose we run `t.hist('y')`. Which histogram does this code produce? `Assign histogram_column_y` to one of the following strings: `'A'`, `'B'`, or `'Neither'`.

_Points:_ 2

```python
histogram_column_y = ...
```

```python
grader.check("task_17")
```

## Uber Movement

According to the [Uber Movement page](https://www.uber.com/us/en/community/supporting-cities/data/):
> Planning great cities requires great data. Uber gathers trip data in more than 10,000 cities across the world. So why not share it? Enter Uber Movement, which gives urban planners access to Uber’s aggregated data to help make informed decisions about our cities.

The Uber Movement project and data access have ended, but we still have data from Boston and Manila during. Below we load tables containing 200,000 weekday Uber rides in the Manila, Philippines, and Boston, Massachusetts metropolitan areas from the Uber Movement project. The `'sourceid'` and `'dstid'` columns contain codes corresponding to start and end locations of each ride. The `'hod'` column contains codes corresponding to the hour of the day the ride took place. The ride time column contains the length of the ride in minutes.

Run the following code cell to create the table `uber` which contains the available Uber ride data for Boston and Manila.

```python
boston = Table.read_table("boston.csv").with_column('city', ['Boston']*200_000)
manila = Table.read_table("manila.csv").with_column('city', ['Manila']*200_000)
uber = boston.append(manila)
uber
```

### Task 18 📍🔎

<!-- BEGIN QUESTION -->

Using the `uber` table, produce an overlaid histogram that visualizes the distributions of all ride times in Boston and Manila. Use the `group='city'` argument with the [`hist` table method](https://datascience.readthedocs.io/en/master/_autosummary/datascience.tables.Table.hist.html#datascience.tables.Table.hist) to accomplish this. Additionally, use the given bins in `equal_bins` by utilizing the `bins` argument for the `hist` table method.

_Points:_ 2

```python
equal_bins = np.arange(0, 120, 5)
...

# Leave the following code to add a title to your histogram
plt.title('Distribution of Boston and Manila Ride Times')
plt.show()
```

<!-- END QUESTION -->

### Task 19 📍🔎

<!-- BEGIN QUESTION -->

Why do you think the distributions for Boston and Manila are different? 

For this task:
* Form a hypothesis that identifies external factors of the two cities that may be causing the difference!
* Provide at least one reference (link) to support your claim.

_Points:_ 2

_Type your answer here, replacing this text._

<!-- END QUESTION -->

### Task 20 📍🔎

<!-- BEGIN QUESTION -->

From the histograms, it looks like there are more 20 to 40-minute Uber rides in Manila compared to Boston. Histograms reflect density, so be careful of interpreting the heights of the bars as counts. It is okay to compare the histograms directly because they both represent 200,000 data points. For this task, produce a bar chart showing that Manila does have more 20 to 40-minute Uber rides than Boston. 

**Note:** For this task, it doesn't matter if you include 40-minute rides in your count or not.

**Hint:** Consider using the [`group` table method](https://datascience.readthedocs.io/en/master/_autosummary/datascience.tables.Table.group.html).

_Points:_ 2

```python
...

# Leave the following code to add a title to your histogram
plt.title('Number of Uber Rides between 20 and 40 Minutes')
plt.show()
```

<!-- END QUESTION -->

## Submit your Homework to Canvas

Once you have finished working on the homework tasks, prepare to submit your work in Canvas by completing the following steps.

1. In the related Canvas Assignment page, check the rubric to know how you will be scored for this assignment.
2. Double-check that you have run the code cell near the end of the notebook that contains the command `"grader.check_all()"`. This command will run all of the run tests on all your responses to the auto-graded tasks marked with 📍.
3. Double-check your responses to the manually graded tasks marked with 📍🔎.
3. Select the menu item "File" and "Save Notebook" in the notebook's Toolbar to save your work and create a specific checkpoint in the notebook's work history.
4. Select the menu items "File", "Download" in the notebook's Toolbar to download the notebook (.ipynb) file. 
5. In the related Canvas Assignment page, click Start Assignment or New Attempt to upload the downloaded .ipynb file.

**Keep in mind that the autograder does not always check for correctness. Sometimes it just checks for the format of your answer, so passing the autograder for a question does not mean you got the answer correct for that question.**

---

To double-check your work, the cell below will rerun all of the autograder tests.

```python
grader.check_all()
```

Tests for task_01 (points: 1)
```
>>> isinstance(unemployment, Table)
True
```
failure_message: ❌ It doesn't seem that you have assigned the name unemployment to a table. Try using the command Table().read_table(...)
success_message: ✅ It seems that you have assigned the name unemployment to a table.
```
>>> np.isclose(np.sum(unemployment.column('NEI')), 1044.46519909577)
True
```
failure_message: ❌ The contents in unemployment don't seem correct. Double-check that you are using the unemployment.csv
success_message: ✅ The contents in unemployment seem correct.


Tests for task_02 (points: 0)
```
>>> isinstance(unemployment.column('DATE').item(0), date)
True
```
failure_message: ❌ It doesn't seem like you've run the cell to convert the date data types. Make sure to run the above code cell.)
success_message: ✅ Great work converting all the dates.


Tests for task_03 (points: 2)
```
>>> by_nei.take(0).column(0).item(0) == date(2020, 4, 1)
True
```
failure_message: ❌ It doesn't look like you have the correct values in by_nei.
success_message: ✅ It seems like you have the correct values in by_nei.
```
>>> by_nei_pter.take(0).column(0).item(0) == date(2020, 4, 1)
True
```
failure_message: ❌ It doesn't look like you have the correct values in by_nei_pter.
success_message: ✅ It seems like you have the correct values in by_nei_pter.


Tests for task_04 (points: 2)
```
>>> greatest_nei.num_rows == 10
True
```
failure_message: ❌ There are not 10 rows (quarters) in your table.
success_message: ✅ There are 10 rows (quarters) in your table.


Tests for task_05 (points: 2)
```
>>> isinstance(pter, np.ndarray)
True
```
failure_message: ❌ pter doesn't seem to be an array.
success_message: ✅ pter is an array.
```
>>> round(pter.item(6), 4) != -1.1446
True
```
failure_message: ❌ Be careful of subtracting the arrays in the wrong order.
success_message: ✅ You didn't seem to subtract the correct arrays in the wrong order.


Tests for task_06 (points: 2)
```
>>> unemployment.num_columns == 4
True
```
failure_message: ❌ Your table doesn't have the correct number of columns.
success_message: ✅ Your table has the correct number of columns.


Tests for task_09 (points: 2)
```
>>> 0 < us_birth_rate < 1
True
```
failure_message: ❌ Make sure you provide a value betwen 0 and 1.
success_message: ✅ You have provided a possible value.


Tests for task_10 (points: 2)
```
>>> 0 < movers <= 52
True
```
failure_message: ❌ Make sure you provide a value betwen 0 and 52 (inclusive).
success_message: ✅ You have provided a possible value.


Tests for task_11 (points: 2)
```
>>> 500000.0 < west_births < 1000000.0
True
```
failure_message: ❌ We are looking for a number between 500000 and 1000000.
success_message: ✅ You have provided a possible value.


Tests for task_12 (points: 2)
```
>>> 0 <= less_than_west_births <= 52
True
```
failure_message: ❌ Make sure you provide a value betwen 0 and 52 (inclusive).
success_message: ✅ You have provided a possible value.


Tests for task_13 (points: 2)
```
>>> visualization in {1, 2, 3}
True
```
failure_message: ❌ Make sure you provide 1, 2, or 3 as a value.
success_message: ✅ You have provided a possible value.


Tests for task_15 (points: 2)
```
>>> type(assoc) is bool
True
```
failure_message: ❌ You didn't provide a possible answer. Use True or False.
success_message: ✅ You provided a possible answer.


Tests for task_16 (points: 2)
```
>>> histogram_column_x in {'A', 'B', 'Neither'}
True
```
failure_message: ❌ Make sure you assign histogram_column_x to either 'A', 'B', or 'Neither'.
success_message: ✅ You provided a possible answer.


Tests for task_17 (points: 2)
```
>>> histogram_column_y in {'A', 'B', 'Neither'}
True
```
failure_message: ❌ Make sure you assign histogram_column_y to either 'A', 'B', or 'Neither'.
success_message: ✅ You provided a possible answer.
