```python
# Initialize Otter
import otter
grader = otter.Notebook("hw09.ipynb")
```

<img style="display: block; margin-left: auto; margin-right: auto" src="./ccsf-logo.png" width="250rem;" alt="The CCSF black and white logo">

# Homework 09: Confidence Intervals

## References

* [Estimation](https://inferentialthinking.com/chapters/13/Estimation.html)
* [`datascience` Documentation](https://datascience.readthedocs.io/)
* [Python Quick Reference](https://ccsf-math-108.github.io/materials-fa23/resources/quick_reference.html)

## Assignment Reminders

- Make sure to run the code cell at the top of this notebook that starts with `# Initialize Otter` to load the auto-grader.
- For all tasks indicated with a 🔎 that you must write explanations and sentences for, provide your answer in the designated space.
- Throughout this assignment and all future ones, please be sure to not re-assign variables throughout the notebook! _For example, if you use `max_temperature` in your answer to one question, do not reassign it later on. Otherwise, you will fail tests that you thought you were passing previously!_
- We encourage you to discuss this assignment with others but make sure to write and submit your own code. Refer to the syllabus to learn more about how to learn cooperatively.
- Unless you are asked otherwise, use the non-interactive visualizations when asked to produce a visualization for a task.
- View the related <a href="https://ccsf.instructure.com" target="_blank">Canvas</a> Assignment page for additional details.

Run the following code cell to import the tools for this assignment.

```python
from datascience import *
import numpy as np
%matplotlib inline
import matplotlib.pyplot as plt
plt.style.use('fivethirtyeight')
```

## Thai Restaurants


Oswaldo and Varun are trying to see what the best Thai restaurant in Berkeley is. They survey 1,500 UC Berkeley students selected uniformly at random and ask each student what Thai restaurant is the best. (*Note: This data is fabricated for the purposes of this homework.*) The choices of Thai restaurants are [Lucky House](https://www.google.com/maps/place/Lucky+House+Thai+Cuisine/@37.8707428,-122.270045,15.32z/data=!4m5!3m4!1s0x80857e9e69a8c921:0x7b6d80f58406fb26!8m2!3d37.8721393!4d-122.2672699), [Imm Thai](https://www.google.com/maps/place/Imm+Thai+Street+Food/@37.8704926,-122.2687372,15.51z/data=!4m5!3m4!1s0x80857e9eec4f1e63:0x5f54d96f0dccdb72!8m2!3d37.8719079!4d-122.2691186), [Thai Temple](https://www.google.com/maps/place/Wat+Mongkolratanaram/@37.8689514,-122.2698649,14.75z/data=!4m5!3m4!1s0x80857e886e39daf1:0xe309caa1b5710fc0!8m2!3d37.8563633!4d-122.2707584), and [Thai Basil](https://www.google.com/maps/place/Thai+Basil/@37.8691911,-122.266539,15.37z/data=!4m5!3m4!1s0x80857c2f6ae0e2f1:0x6978b6e8a72d58d4!8m2!3d37.868327!4d-122.258081). After compiling the results, Oswaldo and Varun release the following percentages, percentages of votes that each restaurant received, from their sample:

|Thai Restaurant  | Percentage|
|:------------:|:------------:|
|Lucky House | 8% |
|Imm Thai | 52% |
|Thai Temple | 25% |
|Thai Basil | 15% |

These percentages represent a uniform random sample of the population of UC Berkeley students. We will attempt to estimate the corresponding *parameters*, or the percentage of the votes that each restaurant will receive from the entire population (the entire population is all UC Berkeley students). We will use confidence intervals to compute a range of values that reflects the uncertainty of our estimates.

The table `votes` contains the results of the survey.

```python
# Just run this cell
votes = Table.read_table('votes.csv').sample(with_replacement = False)
votes
```

### Task 01 📍

Complete the function `one_resampled_percentage` below. It should return Imm Thai's **percentage** of votes after simulating one bootstrap sample of `tbl`.

* You can assume that `tbl` contains the original sample data. Your function should create a resample from this data.
* `tbl` will always be in the same format as `votes`.
* This function should be completed **without** `.group` or `.pivot`. Using these functions will cause your code to timeout.
* Given a table of votes, how can you figure out what percentage of the votes are for a certain restaurant? **Be sure to use percentages, not proportions, for this question!**

_Points:_ 3

```python
def one_resampled_percentage(tbl):
    ...
    
one_resampled_percentage(votes)
```

```python
grader.check("task_01")
```

### Task 02 📍

Complete the `percentages_in_resamples` function such that it simulates and returns an array of **2023 elements**, where each element represents a bootstrapped estimate of the percentage of voters who will vote for Imm Thai. You should use the `one_resampled_percentage` function you wrote above.

* We perform our simulation with **only 2023 trials** in this problem to reduce the runtime, but we should generally use more repetitions.

_Points:_ 3

```python
def percentages_in_resamples():
    percentage_imm = make_array()
    ...
```

In the following cell, we run the function you just defined, `percentages_in_resamples`, and create a histogram of the calculated statistic for the 2023 bootstrap estimates of the percentage of voters who voted for Imm Thai. 

*Note:* This might take a few seconds to run.

```python
resampled_percentages = percentages_in_resamples()
Table().with_column('Estimated Percentage', resampled_percentages).hist("Estimated Percentage")
```

### Task 03 📍

Using the array `resampled_percentages`, find the values at the two edges of the middle 95% of the bootstrapped percentage estimates. (Compute the lower and upper ends of the interval, named `imm_lower_bound` and `imm_upper_bound`, respectively.)


_Points:_ 2

```python
imm_lower_bound = ...
imm_upper_bound = ...
print(f"Bootstrapped 95% confidence interval for the percentage of Imm Thai voters in the population:[{imm_lower_bound:.2f}, {imm_upper_bound:.2f}]")
```

```python
grader.check("task_03")
```

### Task 04 📍

The survey results seem to indicate that Imm Thai is beating all the other Thai restaurants among the voters. We would like to use confidence intervals to determine a range of likely values for Imm Thai's true lead over all the other restaurants combined. The calculation for Imm Thai's lead over Lucky House, Thai Temple, and Thai Basil combined is:

$$\text{Imm Thai's percent of the vote} - (\text{100 percent} - \text{Imm Thai's percent of the vote})$$

Define the function `one_resampled_difference` that returns **exactly one value** of Imm Thai's percentage lead over Lucky House, Thai Temple, and Thai Basil combined from one bootstrap sample of `tbl`.

* Imm Thai's lead can be negative.
* Given a table of votes, how can you figure out what percentage of the votes are for a certain restaurant? **Be sure to use percentages, not proportions, for this question!**
* If the skeleton code provided within the function is not helpful for you, feel free to approach the question using your own variables.


_Points:_ 2

```python
def one_resampled_difference(tbl):
    imm_percentage = ...
    ...
```

```python
grader.check("task_04")
```

### Task 05 📍🔎

<!-- BEGIN QUESTION -->

Write a function called `leads_in_resamples` that returns an array of 2023 elements representing the bootstrapped estimates (the result of calling `one_resampled_difference`) of Imm Thai's lead over Lucky House, Thai Temple, and Thai Basil combined. Afterwards, run the cell to plot a histogram of the resulting samples. 

* If you see an error involving `NoneType`, consider what components a function needs to have!

_Points:_ 2

```python
def leads_in_resamples():
    ...

sampled_leads = leads_in_resamples()
Table().with_column('Estimated Lead', sampled_leads).hist("Estimated Lead")
```

<!-- END QUESTION -->

### Task 06 📍

Use the simulated data in `sampled_leads` from Task 05 to compute an approximate 95% confidence interval for Imm Thai's true lead over Lucky House, Thai Temple, and Thai Basil combined.


_Points:_ 3

```python
diff_lower_bound = ...
diff_upper_bound = ...
print("Bootstrapped 95% confidence interval for Imm Thai's true lead over Lucky House, Thai Temple, and Thai Basil combined: [{:f}, {:f}]".format(diff_lower_bound, diff_upper_bound))
```

```python
grader.check("task_06")
```

## Interpreting Confidence Intervals

We computed the following 95% confidence interval for the percentage of Imm Thai voters: 

$$[49.40, 54.47]$$

(Your answer may have been a bit different; that doesn't mean it was wrong!)

### Task 07 📍🔎

<!-- BEGIN QUESTION -->

Can we say there is a 95% probability that the interval [49.40, 54.47] contains the true percentage of the population that votes for Imm Thai as the best Berkeley Thai restaurant? Answer "yes" or "no" and explain your reasoning. 

*Note:* ambiguous answers using language like "sometimes" or "maybe" will not receive credit.


_Points:_ 2

_Type your answer here, replacing this text._

<!-- END QUESTION -->

### Task 08 📍🔎

<!-- BEGIN QUESTION -->

We also created 70%, 90%, and 99% confidence intervals from the same sample, but we forgot to label which confidence interval represented which percentages! Match each confidence level (70%, 90%, 99%) with its corresponding interval in the cell below (e.g. __ % CI: [49.87, 54.0] $\rightarrow$ replace the blank with one of the three confidence levels). **Then**, explain your thought process and how you came up with your answers.

The intervals are below:

* [49.87, 54.00]
* [50.67, 53.27]
* [48.80, 55.40]


_Points:_ 2

_Type your answer here, replacing this text._

<!-- END QUESTION -->

### Task 09 📍

Suppose we produced 5,000 new samples (each one a uniform random sample of 1,500 voters/students) from the population and created a 95% confidence interval from each one. Roughly how many of those 5,000 intervals do you expect will actually contain the true percentage of the population?

Assign your answer to `true_percentage_intervals`.


_Points:_ 2

```python
true_percentage_intervals = ...
```

```python
grader.check("task_09")
```

Recall the second bootstrap confidence interval you created, which estimated Imm Thai's lead over Lucky House, Thai Temple, and Thai Basil combined. Among
voters in the sample, Imm Thai's lead was 4%. 

Our 95% confidence interval for the true lead (in the population of all voters) was

$$[-0.80, 8.80]$$

Suppose we are interested in testing a simple yes-or-no question:

> "Is the percentage of votes for Imm Thai tied with the percentage of votes for Lucky House, Thai Temple, and Thai Basil combined?"

Our null hypothesis is that the percentages are equal, or equivalently, that Imm Thai's lead is exactly 0. Our alternative hypothesis is that Imm Thai's lead is not equal to 0.  In the questions below, don't compute any confidence interval yourself - use only the 95% confidence interval we created and shared with you.

### Task 10 📍

Say we use a 5% P-value cutoff.  Do we reject the null, fail to reject the null, or are we unable to tell using our confidence interval?

Assign `restaurants_tied` to the number corresponding to the correct answer.

1. Reject the null / Data is consistent with the alternative hypothesis
2. Fail to reject the null / Data is consistent with the null hypothesis
3. Unable to tell using the provided confidence interval

_Consider the relationship between the p-value cutoff and confidence. If you're confused, take a look at [this chapter](https://inferentialthinking.com/chapters/13/4/Using_Confidence_Intervals.html) of the textbook._


_Points:_ 2

```python
restaurants_tied = ...
```

```python
grader.check("task_10")
```

### Task 11 📍

What if, instead, we use a P-value cutoff of 1%? Do we reject the null, fail to reject the null, or are we unable to tell using our confidence interval?

Assign `cutoff_one_percent` to the number corresponding to the correct answer.

1. Reject the null / Data is consistent with the alternative hypothesis
2. Fail to reject the null / Data is consistent with the null hypothesis
3. Unable to tell using the provided confidence interval


_Points:_ 2

```python
cutoff_one_percent = ...
```

```python
grader.check("task_11")
```

### Task 12 📍

What if we use a P-value cutoff of 10%? Do we reject, fail to reject, or are we unable to tell using our confidence interval?

Assign `cutoff_ten_percent` to the number corresponding to the correct answer.

1. Reject the null / Data is consistent with the alternative hypothesis
2. Fail to reject the null / Data is consistent with the null hypothesis
3. Unable to tell using our staff confidence interval


_Points:_ 2

```python
cutoff_ten_percent = ...
```

```python
grader.check("task_12")
```

## Submit your Homework to Canvas

Once you have finished working on the homework tasks, prepare to submit your work in Canvas by completing the following steps.

1. In the related Canvas Assignment page, check the rubric to know how you will be scored for this assignment.
2. Double-check that you have run the code cell near the end of the notebook that contains the command `"grader.check_all()"`. This command will run all of the run tests on all your responses to the auto-graded tasks marked with 📍.
3. Double-check your responses to the manually graded tasks marked with 📍🔎.
3. Select the menu item "File" and "Save Notebook" in the notebook's Toolbar to save your work and create a specific checkpoint in the notebook's work history.
4. Select the menu items "File", "Download" in the notebook's Toolbar to download the notebook (.ipynb) file. 
5. In the related Canvas Assignment page, click Start Assignment or New Attempt to upload the downloaded .ipynb file.

**Keep in mind that the autograder does not always check for correctness. Sometimes it just checks for the format of your answer, so passing the autograder for a question does not mean you got the answer correct for that question.**

---

To double-check your work, the cell below will rerun all of the autograder tests.

```python
grader.check_all()
```

Tests for task_01 (points: 3)
```
>>> isinstance(one_resampled_percentage(votes), (float, np.float_))
True
```
failure_message: ❌ Your function should return a float.
success_message: ✅ Your function returns a float.
```
>>> 35 <= one_resampled_percentage(votes) <= 65
True
```
failure_message: ❌ Your function does not return a value within a range of possible answers.
success_message: ✅ Your function returns a value within a range of possible answers.


Tests for task_02 (points: 3)


Tests for task_03 (points: 2)
```
>>> 40 <= imm_lower_bound <= imm_upper_bound <= 60
True
```
failure_message: ❌ Your values for imm_lower_bound and imm_upper_bound are not within an appropriate interval.
success_message: ✅ Your values for imm_lower_bound and imm_upper_bound are within an appropriate interval.


Tests for task_04 (points: 2)
```
>>> isinstance(one_resampled_difference(votes), (float, np.float_))
True
```
failure_message: ❌ Your function should return a float.
success_message: ✅ Your function returns a float.


Tests for task_06 (points: 3)
```
>>> -5 <= diff_lower_bound <= diff_upper_bound <= 12
True
```
failure_message: ❌ The values of diff_lower_bound and diff_upper_bound are not within a reasonable interval.
success_message: ✅ The values of diff_lower_bound and diff_upper_bound are within a reasonable interval.


Tests for task_09 (points: 2)
```
>>> 1000 <= true_percentage_intervals <= 10000
True
```
failure_message: ❌ The value of true_percentage_intervals is not within a possible interval. 
success_message: ✅ The value of true_percentage_intervals is within a possible interval.


Tests for task_10 (points: 2)
```
>>> 1 <= restaurants_tied <= 3
True
```
failure_message: ❌ restaurants_tied should name 1, 2, or 3.
success_message: ✅ restaurants_tied represents a possible answer.


Tests for task_11 (points: 2)
```
>>> 1 <= cutoff_one_percent <= 3
True
```
failure_message: ❌ cutoff_one_person should name 1, 2, or 3.
success_message: ✅ cutoff_one_percent represents a possible answer.


Tests for task_12 (points: 2)
```
>>> 1 <= cutoff_ten_percent <= 3
True
```
failure_message: ❌ cutoff_ten_percent should name 1, 2, or 3.
success_message: ✅ cutoff_ten_percent represents a possible answer.
