```python
# Initialize Otter
import otter
grader = otter.Notebook("hw05.ipynb")
```

<img style="display: block; margin-left: auto; margin-right: auto" src="./ccsf-logo.png" width="250rem;" alt="The CCSF black and white logo">

# Homework 5: Iteration and Chance

## References

* [Sections 9.0 - 9.3](https://inferentialthinking.com/chapters/09/Randomness.html)
* [Chapter 10](https://inferentialthinking.com/chapters/10/Sampling_and_Empirical_Distributions.html)
* [`datascience` Documentation](https://datascience.readthedocs.io/)
* [Python Quick Reference](https://ccsf-math-108.github.io/materials-fa23/resources/quick_reference.html)

## Assignment Reminders

- Make sure to run the code cell at the top of this notebook that starts with `# Initialize Otter` to load the auto-grader.
- For all tasks indicated with a 🔎 that you must write explanations and sentences for, provide your answer in the designated space.
- Throughout this assignment and all future ones, please be sure to not re-assign variables throughout the notebook! _For example, if you use `max_temperature` in your answer to one question, do not reassign it later on. Otherwise, you will fail tests that you thought you were passing previously!_
- We encourage you to discuss this assignment with others but make sure to write and submit your own code. Refer to the syllabus to learn more about how to learn cooperatively.
- Unless you are asked otherwise, use the non-interactive visualizations when asked to produce a visualization for a task.
- View the related <a href="https://ccsf.instructure.com" target="_blank">Canvas</a> Assignment page for additional details.

Run the following code cell to import the tools for this assignment.

```python
from datascience import *
import numpy as np
%matplotlib inline
import matplotlib.pyplot as plt
plt.style.use('fivethirtyeight')
```

## 2023-24 CCSF Football Season

You are going to analyze how well the CCSF football team performed in the 2023-24 season. A football game is divided into four periods, called quarters. The number of points CCSF scored in each quarter (e.g. `CCSF 1Q`), and the number of points their opponent scored in each quarter  (e.g. `Opp 1Q`) are stored in a table called `games`. If the game is tied at the end of the 4th quarter, the game goes into an additional period called overtime.

**Notes**:
* The 2022-24 season data was collected using the Game Log from [the CCSF Athletics website](https://ccsfathletics.com/sports/fball/2023-24/teams/sanfrancisco?view=gamelog).
* The table `games_ful` contains dditional statistics at the right end of the table for your curiosity.
* A `nan` value indicates that no data was provided on the website.

```python
games_full = Table().read_table("ccsf_fb.csv")
games = games_full.select(1, 3, 4, 5, 6, 7, 8, 9, 10, 11, 12)
games.show()
```

Let's start by finding the total points each team scored in a game.

### Task 01 📍

Write a function called `sum_scores`.  It should take five arguments, where each argument is the team's quarter or overtime score. It should return the team's total score for that game.

_Points:_ 2

```python
def sum_scores(...):
    """Returns the total score calculated by adding up the score of each quarter and the overtime"""
    ...

sum_scores(14, 7, 3, 0, 3) #DO NOT CHANGE THIS LINE
```

```python
grader.check("task_01")
```

### Task 02 📍

Create a new table `final_scores` with three columns in this *specific* order: `Opponent`, `CCSF Score`, `Opponent Score`. You will have to create the `CCSF Score` and `Opponent Score` columns. Use the function `sum_scores` you just defined in the previous question for this problem.

*Hint:* If you want to apply a function that takes in multiple arguments, you can pass multiple column names as arguments in `tbl.apply()`. The column values will be passed into the corresponding arguments of the function. Take a look at the python reference for syntax.

*Tip:* If you’re running into issues creating final_scores, check that `ccsf_scores` and `opponent_scores` output what you want. Also, you will more than likely want to create more steps/intermediate variables.



_Points:_ 3

```python
final_scores = ...
final_scores
```

```python
grader.check("task_02")
```

We can get specific row objects from a table. You can use `tbl.row(n)` to get the `(n+1)`th row of a table. `row.item("column_name")` will allow you to select the element that corresponds to `column_name` in a particular row. Here's an example:

```python
# Just run this cell
games.row(10)
```

```python
# Just run this cell
games.row(10).item("CCSF 4Q")
```

### Task 03 📍

We want to see for a particular game whether or not CCSF won. Write a function called `did_ccsf_win`.  It should take one argument: a **row object** from the `final_scores` table. It should return either `True` if CCSF's score was greater than the Opponent's score, and `False` otherwise.

*Note 1*: "Row object" means a row from the table extracted (behind the scenes) using `tbl.row(index)` that contains all the data for that specific row. It is **not** the index of a row. Do not try and call `final_scores.row(row)` inside of the function.

*Note 2*: If you're still confused by row objects, try printing out `final_scores.row(1)` in a new cell to visually see what it looks like! This piece of code is pulling out the row object located at index 1 of the `final_scores` table and returning it. When you display it in a cell, you'll see that it is not located within a table, but is instead a standalone row object!




_Points:_ 3

```python
def ...(row):
    ...
```

```python
grader.check("task_03")
```

### Task 04 📍

Unfortunately, CCSF did not win against every opponent during the 2023-24 season. Using the `final_scores` table, assign `results` to an array of `True` and `False` values that correspond to whether or not CCSF won. Add the `results` array to the `final_scores` table, and assign this to `final_scores_with_results`. Then, respectively assign the number of wins and losses CCSF had to `ccsf_wins` and `ccsf_losses`.

*Hint*: When you only pass a function name and no column labels through `tbl.apply()`, the function gets applied to every row in `tbl`



_Points:_ 4

```python
results = ...
final_scores_with_results = ...
ccsf_wins = ...
ccsf_losses = ...

# Don't delete or edit the following line:
print(f"In the 2023-24 Season, CCSF Football won {ccsf_wins} games and lost {ccsf_losses} games. Go RAMS!")
```

```python
grader.check("task_04")
```

### Task 05 📍

Sometimes in football, the two teams are equally matched and the game is quite close. Other times, it is a blowout, where the winning team wins by a large margin of victory. Let's define a **big win** to be a game in which the winning team won by more than 10 points. 

Create a function called `is_big_win`.
* The function should accept a single row (`datascience.tables.Row` data type) from a table like the `final_scores` table.
* The function should output `True` if CCSF won by **more than** 10 points. Otherwise, it should return `False`.

Test your function on the first row of the table. `final_scores.row(0)` should be the row object `Row(Opponent='Santa Rosa', CCSF Score=58, Opponent Score=14)`. Since CCSF's score is more than 10 points larger than Santa Rosa's score, then `is_big_win` applied to that row should output `True`.

_Points:_ 4

```python
...

# Test your function 
row1 = final_scores.row(0)
is_big_win(row1)
```

```python
grader.check("task_05")
```

### Task 06 📍

Use your `final_scores` table with your `is_big_win` function to assign `big_wins` to an array of team names that CCSF had big wins against during the 2023-24 football season.

_Points:_ 5

```python
big_wins = ...

for row_index ...
    row = ...
    opponent = ...
    if ...:
        big_wins = np.append(big_wins, ...)

big_wins
```

```python
grader.check("task_06")
```

## Roulette

### Chance

A Nevada roulette wheel has 38 pockets and a small ball that rests on the wheel. When the wheel is spun, the ball comes to rest in one of the 38 pockets. That pocket is declared the winner. 

The pockets are labeled 0, 00, 1, 2, 3, 4, ... , 36. Pockets 0 and 00 are green, and the other pockets are alternately red and black. The table `wheel` is a representation of a Nevada roulette wheel. **Note that *both* columns consist of strings.** Below is an example of a roulette wheel!

<img src="./roulette_wheel.jpeg" alt="roulette wheel" width="330px">

Run the cell below to load the `wheel` table.

```python
wheel = Table.read_table('roulette_wheel.csv', dtype=str)
wheel
```

Before you do the following tasks, make sure you understand the logic behind all the examples in [Section 9.5](https://inferentialthinking.com/chapters/09/5/Finding_Probabilities.html). 

Good ways to approach probability calculations include:

- Thinking one trial at a time: What does the first one have to be? Then what does the next one have to be?
- Breaking up the event into distinct ways in which it can happen.
- Seeing if it is easier to find the chance that the event does not happen.

On each spin of a roulette wheel, all 38 pockets are equally likely to be the winner regardless of the results of other spins. Among the 38 pockets, 18 are red, 18 black, and 2 green.

#### Task 07 📍

The winning pocket is black on all of the first three spins.

_Provide an expression that Python evaluates to the chance of the event described._

_Points:_ 2

```python
first_three_black = ...
```

```python
grader.check("task_07")
```

#### Task 08 📍

The color green never wins in the first 10 spins.

_Provide an expression that Python evaluates to the chance of the event described._

_Points:_ 2

```python
no_green = ...
```

```python
grader.check("task_08")
```

#### Task 09 📍

The color green wins **at least** once on the first 10 spins.

_Provide an expression that Python evaluates to the chance of the event described._

_Points:_ 2

```python
at_least_one_green = ...
```

```python
grader.check("task_09")
```

### Comparing Chances

In each of the following two tasks, two events A and B are described. Choose from one of the following three options and set each answer variable to a single integer:

1. Event A is more likely than Event B
2. Event B is more likely than Event A
3. The two events have the same chance.

You should be able to make the choices **without calculation**. Good ways to approach this exercise include imagining carrying out the chance experiments yourself, one trial at a time, and by thinking about the [law of averages](https://inferentialthinking.com/chapters/10/1/Empirical_Distributions.html#the-law-of-averages).

#### Task 10 📍

A child picks four times at random from a box that has four toy animals: a bear, an elephant, a giraffe, and a kangaroo.

- Event A: all four different animals are picked, assuming the child picks without replacement
- Event B: all four different animals are picked, assuming the child picks with replacement

_Points:_ 2

```python
toys_option = ...
```

```python
grader.check("task_10")
```

#### Task 11 📍

In a lottery, two numbers are drawn at random without replacement from the integers 1 through 1000.

- Event A: The number 8 is picked on both draws
- Event B: The same number is picked on both draws


_Points:_ 2

```python
lottery_option = ...
```

```python
grader.check("task_11")
```

## Submit your Homework to Canvas

Once you have finished working on the homework tasks, prepare to submit your work in Canvas by completing the following steps.

1. In the related Canvas Assignment page, check the rubric to know how you will be scored for this assignment.
2. Double-check that you have run the code cell near the end of the notebook that contains the command `"grader.check_all()"`. This command will run all of the run tests on all your responses to the auto-graded tasks marked with 📍.
3. Double-check your responses to the manually graded tasks marked with 📍🔎.
3. Select the menu item "File" and "Save Notebook" in the notebook's Toolbar to save your work and create a specific checkpoint in the notebook's work history.
4. Select the menu items "File", "Download" in the notebook's Toolbar to download the notebook (.ipynb) file. 
5. In the related Canvas Assignment page, click Start Assignment or New Attempt to upload the downloaded .ipynb file.

**Keep in mind that the autograder does not always check for correctness. Sometimes it just checks for the format of your answer, so passing the autograder for a question does not mean you got the answer correct for that question.**

---

To double-check your work, the cell below will rerun all of the autograder tests.

```python
grader.check_all()
```

Tests for task_01 (points: 2)
```
>>> sum_scores(14, 7, 3, 0, 3) == 27 and sum_scores(2, 3, 6, 1, 0) == 12
True
```
failure_message: ❌ Your function doesn't seem to work correctly with our 4 quarter scores and over time score.
success_message: ✅ Your function works with our 4 quarter scores and over time score.


Tests for task_02 (points: 3)
```
>>> final_scores.num_columns == 3
True
```
failure_message: ❌ final_scores does not have the correct number of columns.
success_message: ✅ final_scores has the correct number of columns.
```
>>> set(['Opponent', 'CCSF Score', 'Opponent Score']) == set(final_scores.labels)
True
```
failure_message: ❌ final_scores does not have the correct labels.
success_message: ✅ final_scores has the correct labels.


Tests for task_03 (points: 3)
```
>>> 'did_ccsf_win' in globals() and callable(did_ccsf_win)
True
```
failure_message: ❌ You didn't create a function called did_ccsf_win.
success_message: ✅ You've created a function called did_ccsf_win.
```
>>> did_ccsf_win(final_scores.row(1)) == True
True
```
failure_message: ❌ Your function doesn't that CCSF won against Sacramento City.
success_message: ✅ Your function shows that CCSF won against Sacramento City.


Tests for task_04 (points: 4)
```
>>> 0 <= ccsf_wins <= 11
True
```
failure_message: ❌ There were only 13 games, so CCSF won 0, 1, 2, ..., or 11 games.
success_message: ✅ ccsf_wins is a possible value
```
>>> 0 <= ccsf_losses <= 11
True
```
failure_message: ❌ There were only 13 games, so CCSF lost 0, 1, 2, ..., or 11 games.
success_message: ✅ ccsf_losses is a possible value


Tests for task_05 (points: 4)
```
>>> callable(is_big_win)
True
```
failure_message: ❌ It doesn't seem like you defined a function called is_big_win.
success_message: ✅ It seems like you defined a function called is_big_win.
```
>>> is_big_win(final_scores.row(0)) == True
True
```
failure_message: ❌ Your function doesn't produce the correct output for the first row of final_scores.
success_message: ✅ Your function produces the correct output for the first row of final_scores.


Tests for task_06 (points: 5)
```
>>> isinstance(big_wins, np.ndarray)
True
```
failure_message: ❌ big_wins should be a NumPy array.
success_message: ✅ big_wins is a NumPy array.
```
>>> type(big_wins.item(0)) == str
True
```
failure_message: ❌ The first item in big_wins is not a string.
success_message: ✅ The first item in big_wins is a string.
```
>>> 0 <= len(big_wins) <= 11
True
```
failure_message: ❌ There is not a possible value for the number items in big_wins.
success_message: ✅ There is a possible value for the number items in big_wins.


Tests for task_07 (points: 2)
```
>>> 0 <= first_three_black <= 1
True
```
failure_message: ❌ first_three_black should be a numerical value between 0 and 1, inclusive.
success_message: ✅ first_three_black is a numerical value between 0 and 1, inclusive.


Tests for task_08 (points: 2)
```
>>> 0 <= no_green <= 1
True
```
failure_message: ❌ no_green should be a numerical value between 0 and 1, inclusive.
success_message: ✅ no_green is a numerical value between 0 and 1, inclusive.


Tests for task_09 (points: 2)
```
>>> 0 <= at_least_one_green <= 1
True
```
failure_message: ❌ at_least_one_green should be a numerical value between 0 and 1, inclusive.
success_message: ✅ at_least_one_green is a numerical value between 0 and 1, inclusive.


Tests for task_10 (points: 2)
```
>>> toys_option in [1, 2, 3]
True
```
failure_message: ❌ 1, 2, or 3 should be assigned to toys_option.
success_message: ✅ This is a possible value for toys_option.


Tests for task_11 (points: 2)
```
>>> lottery_option in [1, 2, 3]
True
```
failure_message: ❌ 1, 2, or 3 should be assigned to lottery_option.
success_message: ✅ This is a possible value for lottery_option.
