```python
# Initialize Otter
import otter
grader = otter.Notebook("project1.ipynb")
```

<img style="display: block; margin-left: auto; margin-right: auto" src="./ccsf-logo.png" width="250rem;" alt="The CCSF black and white logo">

# Project 1: World Population and Poverty

In this project, you'll explore data from [Gapminder.org](http://gapminder.org), a website dedicated to providing a fact-based view of the world and how it has changed. That site includes several data visualizations and presentations, but also publishes the raw data that we will use in this project to recreate and extend some of their most famous visualizations.

The Gapminder website collects data from many sources and compiles them into tables that describe many countries around the world. All of the data they aggregate are published in the [Systema Globalis](https://github.com/open-numbers/ddf--gapminder--systema_globalis/blob/master/README.md). Their goal is "to compile all public statistics; Social, Economic and Environmental; into a comparable total dataset." All data sets in this project are copied directly from the Systema Globalis folder [countries-etc-datapoints](https://github.com/open-numbers/ddf--gapminder--systema_globalis/tree/master/countries-etc-datapoints) on GitHub without any changes. 

*Since the data changes over time, the project will actually use CSV files in the project folder for consistency rather than drawing from the real data that could update while you are working on the project.*

This project is dedicated to [Hans Rosling](https://en.wikipedia.org/wiki/Hans_Rosling) (1948-2017), who championed the use of data to understand and prioritize global development challenges.

## 🎗️ Reminders

* **Canvas Assignment Details**. View the related <a href="https://ccsf.instructure.com" target="_blank">Canvas</a> Assignment page for additional details such as the assignment rubric and due dates.
* **Auto-Grader.** Make sure to run the code cell at the top of this notebook that starts with `# Initialize Otter` to load the auto-grader.
* **[Optional] Project Partner.** 
    * For this assignment, you are allowed to complete this project with a partner, but you are not required to work with a partner.
    * If you want to work with someone, you and your partner should set up a Canvas group with your instructor for this assignment.
    * Keep in mind that for group assignments, one submission will count for the entire group in Canvas.
    * You and your partner will receive the same grade for this project.
* **Code Sharing.** Don't share your code directly with anybody but your partner. You are welcome to discuss questions with other students, but don't share the answers. The experience of solving the problems in this project will prepare you for our exams and potentially for future work in this field. If someone asks you for the answer, resist! Instead, you can demonstrate how you would solve a similar problem or you can focus on a specific part of a task.
* **Support.** You are not alone! Meet with your instructor or the tutors, post questions on Canvas or the student Discord channel, and talk to your classmates during our meetings. If you're ever feeling overwhelmed or don't know how to make progress, talk to your instructor or a tutor.
* **Tests.** The tests that are given are **not comprehensive** and passing the tests for a question **does not** mean that you answered the question correctly. Tests usually only check that your table has the correct column labels. However, more tests will be applied to verify the correctness of your submission in order to assign your final score, so be careful and check your work! You might want to create your own checks along the way to see if your answers make sense. Additionally, before you submit, make sure that none of your cells take a very long time to run (several minutes).
* **Open Response Tasks:** Make sure that you put the answers to the written questions in the indicated cell we provide and generate the graphs as requested. Usually, the point number for the open response tasks corresponds to the number of things you need to address in your response.
* **Advice.** Develop your answers incrementally. To perform a complicated table manipulation, break it up into steps, perform each step on a different line, give a new name to each result, and check that each intermediate result is what you expect. You can add any additional names or functions you want to the provided cells. Make sure that you are using distinct and meaningful variable names throughout the notebook. Along that line, **DO NOT** reuse the variable names that we use when we grade your answers. For example, in the first task, we ask you to assign an answer to `latest`. Do not reassign the variable name `latest` to anything else in your notebook, otherwise there is the chance that our tests grade against what `latest` was reassigned to.
* You are never expected to use just one line in this project or any others. Use intermediate variables and multiple lines as much as you would like! 

Run the following cell to get and import the relevant tools.

```python
from datascience import *
import numpy as np
%matplotlib inline
import matplotlib.pyplot as plt
plt.style.use('fivethirtyeight')
from ipywidgets import interact, interactive, fixed, interact_manual
import ipywidgets as widgets
```

## 📈 Global Population Growth


The global population of humans reached 1 billion around 1800, 3 billion around 1960, 7 billion around 2011, and 8 billion around 2022. The potential impact of population growth has concerned scientists, economists, and politicians alike.

The United Nations Population Division estimates that the world population will likely continue to grow throughout the 21st century, but at a slower rate, perhaps reaching and stabilizing at 11 billion by 2100. However, the UN does not rule out scenarios of slower or more extreme growth. These projections help us understand long-term population processes, even if they leave out possible global catastrophic events like war or climate crises.

The following image was taken from the [United Nation's World Population Prospects 2022 report](https://www.un.org/development/desa/pd/sites/www.un.org.development.desa.pd/files/wpp2022_summary_of_results.pdf).

<img src="UN_WPP_2022_Global_Projection.png" alt="The UN's projection of Global population size and growth rate from 1950 through 2050."/> 

In this part of the project, we will examine some of the factors that influence population growth and how they have been changing over the years and around the world. There are two main sub-parts of this analysis.

- First, we will examine the data for one country, Poland. We will see how factors such as life expectancy, fertility rate, and child mortality have changed over time in Poland, and how they are related to the rate of population growth.
- Next, we will examine whether the changes we have observed for Poland are particular to that country or whether they reflect general patterns observable in other countries too. We will study aspects of world population growth and see how they have been changing.

The first table we will consider contains the total population of each country over time. Run the cell below.

```python
population = Table.read_table('population.csv')
population.show(3)
```

### 🇵🇱 Poland

The Central European nation of Poland has undergone many changes over the centuries. In modern times it was (re)created as a democratic republic in 1919 after World War I. It was invaded and divided in World War II between Germany and the Soviet Union. War and the Holocaust had a devastating impact on its people. Poland was constituted in its current borders at the end of World War II (1945) under a communist government. In 1989, with the fall of the Soviet Union, Poland re-established itself as a democratic republic.

In this section of the project, we will examine aspects of the population of Poland since 1900. Poland's borders have changed, so we will look at the population within its current (2012) borders.

In the `population` table, the `geo` column contains three-letter codes established by the [International Organization for Standardization](https://en.wikipedia.org/wiki/International_Organization_for_Standardization) (ISO) in the [Alpha-3](https://en.wikipedia.org/wiki/ISO_3166-1_alpha-3#Current_codes) standard. Use the Alpha-3 link to find the 3-letter code for Poland.

### Task 01 📍

Create a table called `p_pop` that has two columns labeled `time` and `population_total`. The first column should contain the years from 1900 through 2020 (including both 1900 and 2020) and the second should contain the population of Poland in each of those years.


_Points:_ 4

```python
p_pop = ...
p_pop.show(3)
```

```python
grader.check("task_01")
```

Run the following cell to create a table called `p_five` that has the population of Poland every five years.

```python
p_pop.set_format('population_total', NumberFormatter)

fives = np.arange(1900, 2021, 5) # 1900, 1905, 1910, ...
p_five = p_pop.sort('time').where('time', are.contained_in(fives))
p_five.show(3)
```

Run the following cell to visualize the population over time. Following the devastating effects of World War I and World War II, Poland's population increased steadily from 1950 to 2000 and then leveled off. In the following questions we'll investigate this period of population growth.

```python
p_five.plot(0, 1)
```

### Task 02 📍

Assign `initial` to an array that contains the population for every five year interval from **1900 to 2015** (inclusive). Then, assign `changed` to an array that contains the population for every five year interval from **1905 to 2020** (inclusive). The first array should include both 1900 and 2015, and the second array should include both 1905 and 2020. You should use the `p_five` table to create both arrays, by first filtering the table to only contain the relevant years.

The annual growth rate for a time period is equal to:

$$\left(\left(\frac{\text{Population at end of period}}{\text{Population at start of period}}\right)^{\displaystyle\frac{1}{\text{number of years}}}\right) -1$$

We have provided the code below that uses  `initial` and `changed` in order to add a column to `p_five` called `annual_growth`. **Don't worry about the calculation of the growth rates**; run the test below to test your solution.

If you are interested in how we came up with the formula for growth rates, consult the [growth rates](https://inferentialthinking.com/chapters/03/2/1/Growth.html) section of the textbook.

_Points:_ 6

```python
initial = ...
changed = ...

p_1900_through_2015 = p_five.where('time', are.between_or_equal_to(1900, 2015)) 
p_five_growth = p_1900_through_2015.with_column('annual_growth', (changed/initial)**0.2-1)
p_five_growth.set_format('annual_growth', PercentFormatter)
```

```python
grader.check("task_02")
```

The annual growth rate in Poland has been declining since 1950, as shown in the table below. 

```python
p_five_growth.where('time', are.above_or_equal_to(1950)).show()
```

Next, we'll try to understand what has changed in Poland that might explain the slowing population growth rate. Run the next cell to load three additional tables of measurements about countries over time.

```python
life_expectancy = Table.read_table('life_expectancy.csv').where('time', are.below(2021))
child_mortality = Table.read_table('child_mortality.csv').relabel(2, 'child_mortality_under_5_per_1000_born').where('time', are.below(2021))
fertility = Table.read_table('fertility.csv').where('time', are.below(2021))
```

The `life_expectancy` table contains a statistic that is often used to measure how long people live, called *life expectancy at birth*. This number, for a country in a given year, [does not measure how long babies born in that year are expected to live](http://blogs.worldbank.org/opendata/what-does-life-expectancy-birth-really-mean). Instead, it measures how long someone would live, on average, if the *mortality conditions* in that year persisted throughout their lifetime. These "mortality conditions" describe what fraction of people at each age survived the year. So, it is a way of measuring the proportion of people that are staying alive, aggregated over different age groups in the population.

Run the following cells below to see `life_expectancy`, `child_mortality`, and `fertility`. Refer back to these tables as they will be helpful for answering further questions!

```python
life_expectancy.show(3)
```

```python
child_mortality.show(3)
```

```python
fertility.show(3)
```

### Task 03 📍🔎

<!-- BEGIN QUESTION -->

Is population growing more slowly perhaps because people aren't living as long? Use the `life_expectancy` table to draw a line graph with the years 1950 and later on the horizontal axis that shows how the *life expectancy at birth* has changed in Poland. Make sure you filter the table properly; otherwise, the graph may look strange!


_Points:_ 3

```python
#Fill in code here
...
```

<!-- END QUESTION -->

### Task 04 📍🔎

<!-- BEGIN QUESTION -->

Do the trends in life expectancy in the graph above directly explain why the population growth rate decreased since 1950 in Poland? Why or why not?

_Points:_ 3

_Type your answer here, replacing this text._

<!-- END QUESTION -->

The `fertility` table contains a statistic that is often used to measure how many babies are being born, the *total fertility rate*. This number describes the [number of children a woman would have in her lifetime](https://www.measureevaluation.org/prh/rh_indicators/specific/fertility/total-fertility-rate), on average, if the current rates of birth by age of the mother persisted throughout her child bearing years, assuming she survived through age 49. 

### Task 05 📍

* Complete the function `fertility_over_time`. It takes the Alpha-3 code of a country as `country` and a `start` year. It returns a two-column table with labels `Year` and `Children per woman` with one row per year starting with (and including) the `start` year and including all later years in the `fertility` table. You will probably need to use the `relabeled` method from the `datascience` module to relabel the columns.
* Then, set `poland_code` to the Alpha-3 code for Poland. The provided call `fertility_over_time(poland_code, 1950)` returns a table that shows Poland's annual fertility rates since 1950. 


_Points:_ 5

```python
def fertility_over_time(country, start):
    """Create a two-column table that describes a country's total fertility rate each year."""
    country_fertility = ...
    country_fertility_after_start = ...
    cleaned_table = ...
    ...

poland_code = ...
fertility_over_time(poland_code, 1950).plot(0, 1) # You should *not* change this line.
```

```python
grader.check("task_05")
```

Plotting the fertility rate in Poland since 1950, we see a downward trend.

### Task 06 📍🔎

<!-- BEGIN QUESTION -->

Do the trends in fertility in the graph above help explain why the population growth rate decreased from 1950 to 2020 in Poland? 

_Points:_ 3

_Type your answer here, replacing this text._

<!-- END QUESTION -->

As seen in the [Population Horizons Factsheet No.5]((https://www.ageing.ox.ac.uk/download/143)), lower fertility rates are often associated with lower child mortality rates. We can see if this association is evident in Poland by plotting the relationship between total fertility rate and [child mortality rate per 1000 children](https://en.wikipedia.org/wiki/Child_mortality).

### Task 07 📍

Create a table `poland_since_1950` that contains one row per year starting with 1950 and:

- A column `Year` containing the year
- A column `Children per woman` describing total fertility in Poland that year
- A column `Child deaths per 1000 born` describing child mortality

_Points:_ 8

```python
pol_fertility = ...  # Try starting with the table you built already!
pol_child_mortality = ...
poland_since_1950 =  ...
poland_since_1950
```

```python
grader.check("task_07")
```

Run the following cell to generate a scatter plot from the `poland_since_1950` table you created.

The plot uses **color** to encode data about the `Year` column. The colors, ranging from dark blue to white, represent the passing of time between 1950 and 2020. For example, a point on the scatter plot representing data from the 1950s would appear as **dark blue** and a point from the 2010s would appear as **light blue**.

```python
x_births = poland_since_1950.column("Children per woman")
y_deaths = poland_since_1950.column("Child deaths per 1000 born")
time_colors = poland_since_1950.column("Year")

plt.figure(figsize=(6,6))
plt.scatter(x_births, y_deaths, c=time_colors, cmap="Blues_r")
plt.colorbar()                  
plt.xlabel("Children per woman")
plt.ylabel("Child deaths per 1000 born");
```

### Task 08 📍🔎

<!-- BEGIN QUESTION -->

In one or two sentences, describe the association (if any) that is illustrated by this scatter plot. Does the diagram show any causal relation between fertility and child mortality?

_Points:_ 3

_Type your answer here, replacing this text._

<!-- END QUESTION -->

---

### 👋 Submit your Progress Check to Canvas

1. In the related Canvas Assignment page, check the requirements for a Complete score for this checkpoint.
2. Run the following code cell that will run all of the auto-grader tests for the above tasks.
3. Double-check that you have provided a response to tasks 03, 04, 06, and 08 that will need to be manually reviewed.
4. Select the menu items `File`, `Save and Export Notebook As...`, and `Html_embed` in the notebook's Toolbar to download an HTML version of this notebook file.
5. In the related Canvas Assignment page, click Start Assignment or New Attempt to upload the downloaded HTML file.

**Keep in mind that the auto-grader does not always check for correctness. Sometimes it just checks for the format of your answer, so passing the auto-grader for a question does not mean you got the answer correct for that question.**

```python
try:
    checkpoint_tests = ["task_01", "task_02", "task_05", "task_07"]
    for test in checkpoint_tests:
        display(grader.check(test))
except NameError:
    print('The auto-grader has not be initialized.')
```

---

### 🌐 The World

The changes observed in Poland can also be observed in many other countries: except during periods of extended war, famine, and social chaos, health services generally improve, life expectancy increases, and child mortality decreases. At the same time, the fertility rate often plummets, and where it does, the population growth rate decreases despite increasing longevity.

Run the cell below to generate two overlaid histograms, one for 1962 and one for 2010, that show the distributions of total fertility rates for these two years among all 201 countries in the `fertility` table.

```python
Table().with_columns(
    '1962', fertility.where('time', 1962).column(2),
    '2010', fertility.where('time', 2010).column(2)
).hist(bins=np.arange(0, 10, 0.5), unit='child per woman')
_ = plt.xlabel('Children per woman')
_ = plt.ylabel('Percent per children per woman')
_ = plt.xticks(np.arange(10))
```

### Task 09 📍

Assign `fertility_statements` to an **array** of the numbers of each statement below that can be correctly inferred from these histograms.

1. About the same number (within 3%) of countries had a fertility rate between 3.5 and 4.5 in both 1962 and 2010.
1. In 1962, less than 20% of countries had a fertility rate below 3.
1. At least half of countries had a fertility rate between 5 and 8 in 1962.
1. In 2010, about 40% of countries had a fertility rate between 1.5 and 2.
1. At least half of countries had a fertility rate below 3 in 2010.
1. More countries had a fertility rate above 3 in 1962 than in 2010.


_Points:_ 2

```python
fertility_statements = ...
```

```python
grader.check("task_09")
```

### Task 10 📍🔎

<!-- BEGIN QUESTION -->

Draw a line plot of the world population from 1800 through 2020 (inclusive of both endpoints). The world population is the sum of all of the countries' populations. You should use the `population` table defined earlier in the project. 

_Points:_ 3

```python
#Fill in code here
...
```

<!-- END QUESTION -->

### Task 11 📍

Create a function `stats_for_year` that takes a `year` and returns a table of statistics that addresses the following: 
* The table it returns should have four columns: `geo`, `population_total`, `children_per_woman_total_fertility`, and `child_mortality_under_5_per_1000_born`.
* Each row should contain one unique Alpha-3 country code and three statistics: population, fertility rate, and child mortality for that `year` from the `population`, `fertility` and `child_mortality` tables. 
* Only include rows for which all three statistics are available for the country and year.
* In addition, restrict the result to country codes that appears in `big_50`, an array of the 50 most populous countries in 2020. This restriction will speed up computations later in the project.
* After you write `stats_for_year`, CALL `stats_for_year` on any year between 1960 and 2020. Try to understand the output of stats_for_year.

As you work, consider the following:
* The tests for this question are quite comprehensive, so if you pass the tests, your function is probably correct. However, without calling your function yourself and looking at the output, it will be very difficult to understand any problems you have, so try your best to write the function correctly and check that it works before you rely on the `grader` tests to confirm your work.
* What do all three tables have in common (pay attention to column names)?


_Points:_ 4

```python
# We first create a population table that only includes the 
# 50 countries with the largest 2020 populations. We focus on 
# these 50 countries only so that plotting later will run faster.
big_50 = population.where('time', are.equal_to(2020)).sort("population_total", descending=True).take(np.arange(50)).column('geo')
population_of_big_50 = population.where('time', are.above(1959)).where('geo', are.contained_in(big_50))

def stats_for_year(year):
    """Return a table of the stats for each country that year."""
    p = population_of_big_50.where('time', are.equal_to(year)).drop('time')
    f = fertility.where('time', are.equal_to(year)).drop('time')
    c = child_mortality.where('time', are.equal_to(year)).drop('time')
    ...

# Call stats_for_year for any year between 1960 and 2020.
...
```

```python
grader.check("task_11")
```

### Task 12 📍

Create a table called `pop_by_decade` with two columns called `decade` and `population`, in this order. 
* It has a row for each year that starts a decade, in increasing order starting with 1960 and ending with 2020. For example, 1960 is the start of the 1960's decade. 
* The `population` column contains the total population of all countries included in the result of `stats_for_year(year)` for the first `year` of the decade. You should see that these countries contain most of the world's population.

As you work on this task, consider:
* One approach is to define a function `pop_for_year` that computes this total population, then `apply` it to the `decade` column.  **Think about how you can use the `stats_for_year` function** from the previous question if you want to implement `pop_for_year`. This first test is just a sanity check for your helper function if you choose to use it. **You will not lose points for not implementing the function `pop_for_year`**.
* The cell where you will generate the `pop_by_decade` table is below the cell where you can choose to define the helper function `pop_for_year`. You should define your `pop_by_decade` table in the cell that starts with the table `decades` being defined. 

_Points:_ 0

```python
def pop_for_year(year):
    ...
```

```python
grader.check("task_12_pop_for_year_check")
```

Now that you've defined your helper function (if you've chosen to do so), define the `pop_by_decade` table.


_Points:_ 7

```python
decades = Table().with_column('decade', np.arange(1960, 2021, 10))

pop_by_decade = ...
pop_by_decade.set_format(1, NumberFormatter)
```

```python
grader.check("task_12")
```

The `countries` table describes various characteristics of countries. The `country` column contains the same codes as the `geo` column in each of the other data tables (`population`, `fertility`, and `child_mortality`). The `world_6region` column classifies each country into a region of the world. Run the cell below to inspect the data.

```python
countries = Table.read_table('countries.csv').where('country', are.contained_in(population.group('geo').column('geo')))
countries.select('country', 'name', 'world_6region')
```

### Task 13 📍

Create a table called `region_counts`. 
* It should contain two columns called `region` and `count`. 
    * The `region` column should contain regions of the world
    * The`count` column should contain the number of countries in each region that appears in the result of `stats_for_year(2020)`. 

For example, one row would have `south_asia` as its `region` value and an integer as its `count` value: the number of large South Asian countries for which we have population, fertility, and child mortality numbers from 2020.

As you work, consider that you may have to relabel a column to name it `region`.

_Points:_ 4

```python
stats_for_2020 = ...
region_counts = ...
region_counts
```

```python
grader.check("task_13")
```

The following scatter diagram compares total fertility rate and child mortality rate for each country in 1960. The area of each dot represents the population of the country, and the color represents its region of the world. Run the cell. Do you think you can identify any of the dots?

```python
from functools import lru_cache as cache
import pandas as pd
import plotly.express as px

# This cache annotation makes sure that if the same year
# is passed as an argument twice, the work of computing
# the result is only carried out once.
@cache(None)
def stats_relabeled(year):
    """Relabeled and cached version of stats_for_year."""
    return stats_for_year(year).relabel(2, 'Children per woman').relabel(3, 'Child deaths per 1000 born')

def fertility_vs_child_mortality(year):
    """Draw a color scatter diagram comparing child mortality and fertility."""
    with_region = stats_relabeled(year).join('geo', countries.select('country', 'world_6region'), 'country') 
    fig = px.scatter(with_region.to_df(),
                     'Children per woman', 
                     'Child deaths per 1000 born',
                     size='population_total',
                     color='world_6region',
                     title=year)
    fig.update_xaxes(range=[0, 10])
    fig.update_yaxes(range=[-50, 500])
    fig.update_layout(width=1000, height=600)
    fig.update_traces(marker=dict(sizemode='area', sizeref=30_000))
    fig.show()
    
fertility_vs_child_mortality(1960)
```

### Task 14 📍

Assign `scatter_statements` to an array of the numbers of each statement below that can be inferred from this scatter diagram for 1960. 
1. All countries in `europe_central_asia` had uniformly low fertility rates.
1. The lowest child mortality rate of any country was from an `east_asia_pacific` country.
1. Most countries had a fertility rate above 5.
1. There was an association between child mortality and fertility.
1. The two largest countries by population also had the two highest child mortality rates.

_Points:_ 2

```python
scatter_statements = ...
```

```python
grader.check("task_14")
```

The result of the cell below is interactive. Drag the slider to the right to see how countries have changed over time. You’ll find that in terms of population growth, the divide between the countries of the global North and global South that existed in the 1960s has shrunk significantly. 

This shift in fertility rates is the reason that the global population is expected to grow more slowly in the 21st century than it did in the 19th and 20th centuries. Fertility rates change for reasons that include cultural patterns, better prospects for children surviving to adulthood, and family planning (such as contraception and women’s greater control over their reproduction).

**Note:** Don't worry if a red warning pops up when running the cell below. You'll still be able to run the cell!

```python
_ = widgets.interact(fertility_vs_child_mortality, year=widgets.IntSlider(min=1960, max=2020, value=1960))
```

Now is a great time to take a break and watch the same data presented by [Hans Rosling in a 2010 TEDx talk](https://www.gapminder.org/videos/reducing-child-mortality-a-moral-and-environmental-imperative) with smoother animation and witty commentary.

```python
from IPython.display import YouTubeVideo
YouTubeVideo("N-x7eHuUhNM")
```

When we look at population and fertility as data scientists, we need to learn about the experiences of people in real life, not just abstractly as data. We should also recognize that population studies have sometimes had political undercurrents. Those undercurrents have included population control, control of reproduction, or fears of shifts between racial groups. To do better as data scientists, we should check our assumptions to avoid unthinkingly reproducing past patterns.

---

## 💰 Global Poverty


In 1800, 85% of the world's 1 billion people lived in [extreme poverty](https://en.wikipedia.org/wiki/Extreme_poverty), defined by the United Nations as "a condition characterized by severe deprivation of basic human needs, including food, safe drinking water, sanitation facilities, health, shelter, education and information." At the time when the data in this project were gathered, a common definition of extreme poverty was a person living on less than \$1.90 per day.

In 2018, the proportion of people living in extreme poverty was estimated to be [about 9\%](https://unstats.un.org/sdgs/report/2019/goal-01/). Although the world rate of extreme poverty has declined consistently for hundreds of years, the number of people living in extreme poverty is still over 600 million. The United Nations adopted an [ambitious goal](http://www.un.org/sustainabledevelopment/poverty/): "By 2030, eradicate extreme poverty for all people everywhere." However, recently, researchers projected on [the World Bank blog](https://blogs.worldbank.org/opendata/updated-estimates-impact-covid-19-global-poverty-turning-corner-pandemic-2021) that there will be an increase in extreme poverty since the start of COVID-19 pandemic.

In this part of the project, we will examine some aspects of global poverty that might affect whether the goal is achievable. The causes of poverty are complex. They include global histories, such as colonialism, as well as factors such as health care, economics, and social inequality in each country.

First, load the population and poverty rate by country and year and the country descriptions. While the `population` table has values for every recent year for many countries, the `poverty` table only includes certain years for each country in which a measurement of the rate of extreme poverty was available.

```python
population = Table.read_table('population.csv')
countries = Table.read_table('countries.csv').where('country', are.contained_in(population.group('geo').column('geo')))
poverty = Table.read_table('poverty.csv')
poverty.show(3)
```

### Task 15 📍

Assign `latest_poverty` to a three-column table with one row for each country that appears in the `poverty` table. The first column should contain the 3-letter code for the country. The second column should contain the most recent year for which an extreme poverty rate is available for the country. The third column should contain the poverty rate in that year. **Do not change the last line, so that the labels of your table are set correctly.**

*Hint*: think about how ```group``` works: it does a sequential search of the table (from top to bottom) and collects values in the array in the order in which they appear, and then applies a function to that array. The `first` function may be helpful, but you are not required to use it.


_Points:_ 3

```python
def first(values):
    return values.item(0)

latest_poverty = ...
latest_poverty = latest_poverty.relabeled(0, 'geo').relabeled(1, 'time').relabeled(2, 'poverty_percent') # You should *not* change this line.  
latest_poverty
```

```python
grader.check("task_15")
```

### Task 16 📍

Using both `latest_poverty` and `population`, create a four-column table called `recent_poverty_total` with one row for each country in `latest_poverty`. The four columns should have the following labels and contents:
1. `geo` contains the 3-letter country code,
1. `poverty_percent` contains the most recent poverty percent,
1. `population_total` contains the population of the country in 2010,
1. `poverty_total` contains the number of people in poverty **rounded to the nearest integer**, based on the 2010 population and most recent poverty rate.

You are **not** required to use `poverty_and_pop`, and you are always welcome to add any additional names.

_Points:_ 4

```python
poverty_and_pop = ...
recent_poverty_total = ...
recent_poverty_total
```

```python
grader.check("task_16")
```

### Task 17 📍

Assign the name `poverty_percent` to the known percentage of the world’s 2010 population that were living in extreme poverty. Assume that the `poverty_total` numbers in the `recent_poverty_total` table describe **all** people in 2010 living in extreme poverty. You should find a number that is above the 2018 global estimate of 8%, since many country-specific poverty rates are older than 2018.

*Hint*: The sum of the `population_total` column in the `recent_poverty_total` table is not the world population, because only a subset of the world's countries are included in the `recent_poverty_total` table (only some countries have known poverty rates). Use the `population` table to compute the world's 2010 total population.


_Points:_ 2

```python
poverty_percent = ...
poverty_percent
```

```python
grader.check("task_17")
```

The `countries` table includes not only the name and region of countries, but also their positions on the globe.

```python
countries.select('country', 'name', 'world_4region', 'latitude', 'longitude')
```

### Task 18 📍

Using both `countries` and `recent_poverty_total`, create a five-column table called `poverty_map` with one row for every country in `recent_poverty_total`.  The five columns should have the following labels and contents, in this order:
1. `latitude` contains the country's latitude,
1. `longitude` contains the country's longitude,
1. `name` contains the country's name,
1. `region` contains the country's region from the `world_4region` column of `countries`,
1. `poverty_total` contains the country's poverty total.


_Points:_ 3

```python
poverty_map = ...
poverty_map
```

```python
grader.check("task_18")
```

Run the cell below to draw a map of the world in which the areas of circles represent the number of people living in extreme poverty according to this data. Double-click on the map to zoom in.

```python
# It may take a few seconds to generate this map.
colors = {'africa': 'blue', 'europe': 'black', 'asia': 'red', 'americas': 'green'}
scaled = poverty_map.with_columns(
    'labels', poverty_map.column('poverty_total'),
    'colors', poverty_map.apply(colors.get, 'region'),
    'areas', 1e-4 * poverty_map.column('poverty_total') + 1
).drop('name', 'region', 'poverty_total')

Circle.map_table(scaled)
```

Although people live in extreme poverty throughout the world (with more than 5 million in the United States), this data shows the largest numbers are in Asia and Africa. Keep in mind that what you are observing is based on the data that was collected. There are clear inconsistencies with what is seen in the data and what people actually experience in the world.

### Task 19 📍

Assign `largest` to a two-column table with the `name` (not the 3-letter code) and `poverty_total` of the 10 countries with the largest number of people living in extreme poverty.


_Points:_ 4

```python
largest = ...
largest.set_format('poverty_total', NumberFormatter)
```

```python
grader.check("task_19")
```

### Task 20 📍🔎

<!-- BEGIN QUESTION -->

It is important to study the absolute number of people living in poverty, not just the percent. The absolute number is an important factor in determining the amount of resources needed to support people living in poverty. In the next two questions you will explore this.

In the task after this, you will be asked to write a function called `poverty_timeline` that takes the name of a country as its argument (not the Alpha-3 country code). It should draw a line plot of the number of people living in poverty in that country with time on the horizontal axis. The line plot should have a point for each row in the poverty table for that country. To compute the population living in poverty from a poverty percentage, multiply by the population of the country in that year.

For this task, write out a generalized process for the function you will write in the next task. What should this function output, and what steps will you take within the function? Keep in mind you are not focusing on writing the code, you are focusing on explaining the concept of how the function should work.

_Points:_ 3

_Type your answer here, replacing this text._

<!-- END QUESTION -->

### Task 21 📍🔎

<!-- BEGIN QUESTION -->

Now, we'll actually write the function called `poverty_timeline`. Recall that `poverty_timeline` takes **the name of a country** as its argument (not the Alpha-3 country code). 
* It should draw a line plot of the number of people living in poverty in that country with time on the horizontal axis. 
* The line plot should have a point for each row in the `poverty` table for that country. To compute the population living in poverty from a poverty percentage, multiply by the population of the country **in that year**.

Consider the following while working on this task:
* This question is long. Feel free to create cells and experiment. You can create cells by going to the toolbar and hitting the `+` button.
* It might be helpful to join the information from different tables in your code.

_Points:_ 3

```python
def poverty_timeline(country):
    '''Draw a timeline of people living in extreme poverty in a country.'''
    geo = ...
    # This solution will take multiple lines of code. Use as many as you need
    ...
    
    # Don't change anything below this line. 
    plt.title(country)
    plt.ylim(bottom=0)
    plt.show() # This should be the last line of your function. 
...
```

<!-- END QUESTION -->

Finally, draw the timeline below for India. You can check your work by comparing your graph to the one on [gapminder.org](https://www.gapminder.org/tools/#$state$entities$show$country$/$in@=ind;;;;&marker$axis_y$which=number_of_people_in_poverty&scaleType=linear&spaceRef:null;;;&chart-type=linechart).


```python
poverty_timeline('Nigeria')
poverty_timeline('United States')
poverty_timeline('China')
poverty_timeline('India')
```

### Task 22 📍🔎

<!-- BEGIN QUESTION -->

According to the data, the number of people living in extreme poverty has been increasing in Nigeria and the United States, and the massive decreases in China and India have shaped the overall trend that extreme poverty is decreasing worldwide, both in percentage and in absolute number. 

To learn more, watch [Hans Rosling in a 2015 film](https://www.gapminder.org/videos/dont-panic-end-poverty/) about the UN goal of eradicating extreme poverty from the world. 

Below, we've also added an interactive dropdown menu for you to visualize `poverty_timeline` graphs for other countries. Note that each dropdown menu selection may take a few seconds to run.

For this task, make sure to address the following. Your response should be about 3-4 sentences, but you are welcome to provide more details.

* Select a country that is not Nigeria, the United States, China, or India in the dropdown menu of the widget below to create the plot showing poverty numbers over the time periods available in the data set.
* In general, describe how the number of those in extreme poverty in that particular country has changed over time.
* Almost every graph has some interesting point(s) in the line where there was a change in the trend of the poverty numbers. Connect the pattern you see to at least one historical observation about that particular country or global event. Include a link to a resource that provides context for that historical observation. Sometimes these changes are associated with changes in political leadership, war, famine, economic development, etc.

_Points:_ 3

```python
all_countries = poverty_map.column('name')
_ = widgets.interact(poverty_timeline, country=list(all_countries))
```

_Type your answer here, replacing this text._

<!-- END QUESTION -->

**You're finished!** Congratulations on working through your first project on data visualization and table manipulation.

## 🏁 Submit your Project to Canvas

Once you have finished working on the projects tasks, prepare to submit your work in Canvas by completing the following steps.

1. In the related Canvas Assignment page, check the rubric to know how you will be scored for this assignment.
2. Double-check that you have run the code cell near the end of the notebook that contains the command `"grader.check_all()"`. This command will run all of the run tests on all your responses to the auto-graded tasks marked with 📍.
3. Double-check your responses to the manually graded tasks marked with 📍🔎.
3. Select the menu item "File" and "Save Notebook" in the notebook's Toolbar to save your work and create a specific checkpoint in the notebook's work history.
4. Select the menu items "File", "Download" in the notebook's Toolbar to download the notebook (.ipynb) file. 
5. In the related Canvas Assignment page, click Start Assignment or New Attempt to upload the downloaded .ipynb file.
    * _Keep in mind that if you are working in a group, then you only submit one file for your group._

**Keep in mind that the auto-grader does not always check for correctness. Sometimes it just checks for the format of your answer, so passing the autograder for a question does not mean you got the answer correct for that question.**

---

To double-check your work, the cell below will rerun all of the autograder tests.

```python
grader.check_all()
```

Tests for task_01 (points: 4)
```
>>> p_pop.labels == ('time', 'population_total')
True
```
failure_message: ❌ Check your column labels and spelling for p_pop
success_message: ✅ It seems like you have the correct columns in p_pop.
```
>>> all(p_pop.sort('time').column('time') == np.arange(1900, 2021))
True
```
failure_message: ❌ The values for time should range from 1900 through 2020, inclusive.
success_message: ✅ It seems the you have the correct time values in p_pop


Tests for task_02 (points: 6)
```
>>> first = round(p_five_growth.sort(0).column(2).item(0), 8)
>>> 0.005 <= first <= 0.5
True
```
failure_message: ❌ We checked one of the values in p_five_growth, and it should be between 0.005 and 0.5
success_message: ✅ One of the values we checked in p_five_growth is between 0.005 and 0.05.
```
>>> max(p_five_growth.column(2)) < 0.02
True
```
failure_message: ❌ Your largest growth rate should be less than 0.02.
success_message: ✅ It seems that your largest growth rate is less than 0.02.


Tests for task_05 (points: 5)
```
>>> fertility_over_time('usa', 2010).labels == ('Year', 'Children per woman')
True
```
failure_message: ❌ Your column labels are not correct. You might need to double check your spelling.
success_message: ✅ You seem to have the correct column labels!
```
>>> all(fertility_over_time('usa', 2010).column('Year') == np.arange(2010, 2021))
True
```
failure_message: ❌ It seems like you are missing some years in the Year column. Make sure to use the start year to determine the data range.
success_message: ✅ Your years seem to span the correct interval.
```
>>> all(fertility_over_time('usa', 2005).column('Year') == np.arange(2005, 2021))
True
```
failure_message: ❌ It seems like you are missing some years in the Year column. Make sure to use the start year to determine the data range.
success_message: ✅ Your years seem to span the correct interval.


Tests for task_07 (points: 8)
```
>>> poland_since_1950.num_rows == 71
True
```
failure_message: ❌ Make sure you are using the date range 1950 - 2020, including 1950 and 2020.
success_message: ✅ It looks like you've covered the correct range of years.
```
>>> all([label in poland_since_1950.labels for label in ['Children per woman', 'Child deaths per 1000 born']])
True
```
failure_message: ❌ Your labels in poland_since_1950 are not correct.
success_message: ✅ Your labels in poland_since_1950 seem correct.


Tests for task_09 (points: 2)
```
>>> import numpy as np
>>> isinstance(fertility_statements, np.ndarray) and all((x in range(1, 7) for x in set(fertility_statements)))
True
```
failure_message: ❌ Make sure you are using make_array with numbers between 1 and 6, inclusive.
success_message: ✅ Great job formatting your answer as an array of integers.


Tests for task_11 (points: 4)
```
>>> t = stats_for_year(1990)
>>> t.labels == ('geo', 'population_total', 'children_per_woman_total_fertility', 'child_mortality_under_5_per_1000_born')
True
```
failure_message: ❌ Your function generates a table, but the columns are not correct.
success_message: ✅ Your function generates a table with the correct columns.
```
>>> t = stats_for_year(1990)
>>> t.num_rows == 50
True
```
failure_message: ❌ Your function produces a table, but the number of rows are not correct.
success_message: ✅ Your function produces a table with the correct number of rows.
```
>>> print(sum(stats_for_year(1960).sort('geo').take(np.arange(5, 50, 5)).column(2)))
49.47
```
failure_message: ❌ Your function didn't seem to create the right fertility values for the year 1960.
success_message: ✅ It seems like your function created the right fertility values for the year 1960.
```
>>> print(sum(stats_for_year(2010).sort('geo').take(np.arange(3, 50, 5)).column(3)))
533.54
```
failure_message: ❌ Your function didn't seem to create the right child mortality values for the year 1960.
success_message: ✅ It seems like your function created the right child mortality values for the year 1960.


Tests for task_12 (points: 7)
```
>>> pop_by_decade.labels == ('decade', 'population')
True
```
failure_message: ❌ pop_by_decade doesn't have the correct labels.
success_message: ✅ pop_by_decade has the correct labels.
```
>>> pop_by_decade.column(0).item(0) == 1960
True
```
failure_message: ❌ The first year of the 1960s should be 1960.
success_message: ✅ The first year of the 1960s is 1960.


Tests for task_12_pop_for_year_check (points: 0)
```
>>> pop_for_year(1972) == 3355562066
True
```
failure_message: ❌ The function doesn't seem to work correctly for 1972.
success_message: ✅ The function seems to work correctly for 1972.
```
>>> pop_for_year(1989) == 4576679948
True
```
failure_message: ❌ The function doesn't seem to work correctly for 1989.
success_message: ✅ The function seems to work correctly for 1989.
```
>>> pop_for_year(2020) == 5514129422
False
```
failure_message: ❌ The function doesn't seem to work correctly for 2002.
success_message: ✅ The function seems to work correctly for 2002.


Tests for task_13 (points: 4)
```
>>> region_counts.labels == ('region', 'count')
True
```
failure_message: ❌ region_counts doesn't contain the correct labels.
success_message: ✅ region_counts contains the correct labels.
```
>>> sum(region_counts.column('count')) == 50
True
```
failure_message: ❌ The count sum should total 50.
success_message: ✅ The count sum totals 50 as expected.


Tests for task_14 (points: None)
```
>>> import numpy as np
>>> isinstance(scatter_statements, np.ndarray) and all((x in range(1, 6) for x in set(scatter_statements)))
True
```
failure_message: ❌ You haven't provided an array of possible values.
success_message: ✅ You have provided an array of possible values.


Tests for task_15 (points: 3)
```
>>> latest_poverty.labels == ('geo', 'time', 'poverty_percent')
True
```
failure_message: ❌ The labels of latest_poverty are not correct.
success_message: ✅ The labels of latest_poverty are correct.
```
>>> latest_poverty.num_rows == 168
True
```
failure_message: ❌ latest_poverty does not have the correct number of rows.
success_message: ✅ latest_poverty has the correct number of rows.


Tests for task_16 (points: 4)
```
>>> recent_poverty_total.labels == ('geo', 'poverty_percent', 'population_total', 'poverty_total')
True
```
failure_message: ❌ recent_poverty_total does not have the correct labels.
success_message: ✅ recent_poverty_total has the correct labels.
```
>>> recent_poverty_total.where('geo', 'aus').column(2).item(0) == 22154687
True
```
failure_message: ❌ recent_poverty_total should show that Australia's population in 2010 is 22,154,687.
success_message: ✅ recent_poverty_total shows correctly that Australia's population in 2010 is 22,154,687.
```
>>> float(recent_poverty_total.where('geo', 'aus').column(3).item(0)) == 110773
True
```
failure_message: ❌ The number of people estimated to be living in extreme poverty in Australia should be 110,773.
success_message: ✅ The number of people estimated to be living in extreme poverty in Australia is 110,773.


Tests for task_17 (points: 2)
```
>>> 10 <= poverty_percent <= 20
True
```
failure_message: ❌ poverty_precent should be between 10 and 20
success_message: ✅ poverty_percent is between 10 and 20.


Tests for task_18 (points: 3)
```
>>> poverty_map.labels == ('latitude', 'longitude', 'name', 'region', 'poverty_total')
True
```
failure_message: ❌ poverty_map does not have the correct labels.
success_message: ✅ poverty_map has the correct labels.
```
>>> list(np.sort(np.unique(poverty_map.column('region'))))
['africa', 'americas', 'asia', 'europe']
```
failure_message: ❌ poverty_map does not have the correct region values.
success_message: ✅ poverty_map has the correct region values.


Tests for task_19 (points: 4)
```
>>> largest.labels == ('name', 'poverty_total')
True
```
failure_message: ❌ largest does not contain the correct labels.
success_message: ✅ largest contains the correct labels.
```
>>> largest.column(0).item(0) == 'India'
True
```
failure_message: ❌ According to the data, India is the country with the largest number of people in extreme poverty.
success_message: ✅ largest correctly identifies (according to this data) India as the country with the largest number of people in extreme poverty.
```
>>> largest.num_rows == 10
True
```
failure_message: ❌ largest should contain 10 rows.
success_message: ✅ largest contains 10 rows.
